[1680764418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.428 second response time
[1680764423] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680764506] SERVICE ALERT: vswprdawbws11;serGEN vswprd Airweb - db2 established connections;OK;HARD;3;OK - 4 established connections to port 50001
[1680764663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 601 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6181 MiB (67.79% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 1842829 MiB (39.05% inode=100%):
[1680764719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.462 second response time
[1680764869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 8 Hours, 42 Minutes
[1680764962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 601 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 1791327 MiB (37.96% inode=100%):
[1680765018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.367 second response time
[1680765061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 371 Days, 23 Hours, 25 Minutes
[1680765115] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765180] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680765244] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 601 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 1736224 MiB (36.79% inode=100%):
[1680765309] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680765311] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.460 second response time
[1680765375] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680765498] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765555] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 601 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 1677542 MiB (35.55% inode=100%):
[1680765562] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680765614] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.426 second response time
[1680765620] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680765678] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680765752] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 8 Hours, 57 Minutes
[1680765817] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680765862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 601 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 1631500 MiB (34.57% inode=100%):
[1680765881] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680765919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.410 second response time
[1680765951] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680765961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 371 Days, 23 Hours, 40 Minutes
[1680765978] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766016] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680766042] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680766076] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766081] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766107] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766141] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680766145] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680766162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 601 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4588 MiB (89.79% inode=100%): /nvol_prd_db2_bku_01 1602342 MiB (33.95% inode=100%):
[1680766171] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680766218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.371 second response time
[1680766236] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680766343] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766407] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680766462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 600 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 1579245 MiB (33.46% inode=100%):
[1680766472] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680766505] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.425 second response time
[1680766523] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766541] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680766570] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680766573] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766587] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680766630] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766634] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680766638] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680766652] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680766670] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 9 Hours, 12 Minutes
[1680766694] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680766702] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680766762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 600 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 1557362 MiB (33.00% inode=100%):
[1680766819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.426 second response time
[1680766862] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 371 Days, 23 Hours, 55 Minutes
[1680766885] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680766950] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 600 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 1537700 MiB (32.58% inode=100%):
[1680767118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.489 second response time
[1680767127] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767192] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767256] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767314] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767321] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 600 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 1518293 MiB (32.17% inode=100%):
[1680767379] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.389 second response time
[1680767445] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767510] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 9 Hours, 27 Minutes
[1680767574] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767639] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680767662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 568 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 1498236 MiB (31.75% inode=100%):
[1680767669] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767703] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680767718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.435 second response time
[1680767734] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767744] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 10 Hours, 10 Minutes
[1680767798] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767808] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680767840] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767863] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767873] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680767904] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680767927] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680767962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 568 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 1478090 MiB (31.32% inode=100%):
[1680767992] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680768008] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680768019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.524 second response time
[1680768057] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680768157] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768178] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680768178] SERVICE ALERT: vswprdawbwa12;serGEN vswprd RHEL Load;WARNING;SOFT;1;WARNING - load average: 3.04, 1.25, 0.59
[1680768222] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680768238] SERVICE ALERT: vswprdawbwa12;serGEN vswprd RHEL Load;OK;SOFT;2;OK - load average: 1.25, 1.06, 0.57
[1680768262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 568 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13618 MiB (88.71% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.31% inode=100%): /var/log/audit 4581 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 1463375 MiB (31.01% inode=100%):
[1680768285] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.514 second response time
[1680768350] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680768414] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 9 Hours, 42 Minutes
[1680768479] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680768562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 568 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 1447862 MiB (30.68% inode=100%):
[1680768576] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768586] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768587] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.418 second response time
[1680768630] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768641] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680768650] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680768652] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680768661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 25 Hours, 25 Minutes
[1680768661] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768695] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680768706] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768715] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680768716] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680768726] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680768744] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768759] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680768770] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680768791] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680768806] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768808] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680768835] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680768858] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 568 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 1425393 MiB (30.20% inode=100%):
[1680768871] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680768909] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680768918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.239 second response time
[1680768923] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680768935] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680768974] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680769020] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769039] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680769043] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680769044] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680769063] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680769085] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680769103] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578222 bytes in 0.210 second response time
[1680769104] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578258 bytes in 0.188 second response time
[1680769139] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680769150] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680769162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 567 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 1406728 MiB (29.81% inode=100%):
[1680769218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.440 second response time
[1680769240] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680769369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 9 Hours, 57 Minutes
[1680769443] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 567 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 1390975 MiB (29.47% inode=100%):
[1680769468] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769508] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.263 second response time
[1680769533] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 40 Hours, 40 Minutes
[1680769587] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769598] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769621] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769639] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769644] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769652] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680769663] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680769683] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769685] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769704] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769708] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769717] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680769727] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680769743] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769748] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2018 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 1382878 MiB (29.30% inode=100%):
[1680769808] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680769812] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769812] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.427 second response time
[1680769873] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680769877] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769877] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680769942] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680769968] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680770006] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680770021] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680770032] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680770062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2007 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.31% inode=100%): /var/log/audit 4584 MiB (89.72% inode=100%): /nvol_prd_db2_bku_01 1375580 MiB (29.15% inode=100%):
[1680770097] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680770118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.295 second response time
[1680770177] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680770269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 10 Hours, 12 Minutes
[1680770296] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680770332] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680770361] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680770363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2007 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.31% inode=100%): /var/log/audit 4584 MiB (89.72% inode=100%): /nvol_prd_db2_bku_01 1367812 MiB (28.98% inode=100%):
[1680770370] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680770373] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680770397] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680770402] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680770419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.280 second response time
[1680770428] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680770435] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680770438] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680770461] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680770462] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 55 Hours, 55 Minutes
[1680770493] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680770502] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680770626] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680770662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2007 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4581 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 1360642 MiB (28.83% inode=100%):
[1680770691] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680770718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.269 second response time
[1680770962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2007 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4581 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 1352440 MiB (28.66% inode=100%):
[1680771007] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.322 second response time
[1680771026] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771044] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771071] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771090] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771100] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771109] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680771141] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771155] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771157] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771164] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 10 Hours, 27 Minutes
[1680771173] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680771206] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771220] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680771222] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771255] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2007 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4581 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 1343628 MiB (28.47% inode=100%):
[1680771284] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680771319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.232 second response time
[1680771319] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771356] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 1 Hours, 10 Minutes
[1680771372] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771384] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771420] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771437] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680771448] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771478] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771485] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771501] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680771513] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771542] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680771550] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2007 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4581 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 1334724 MiB (28.28% inode=100%):
[1680771577] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771589] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680771607] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680771618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.575 second response time
[1680771654] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680771806] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680771862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2016 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4581 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 1326467 MiB (28.11% inode=100%):
[1680771919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.343 second response time
[1680772019] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 10 Hours, 42 Minutes
[1680772084] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680772111] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2016 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4581 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 1318781 MiB (27.94% inode=100%):
[1680772170] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772175] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680772218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.268 second response time
[1680772235] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680772240] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 1 Hours, 25 Minutes
[1680772304] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680772369] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680772462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2016 MiB (0.98% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4577 MiB (89.58% inode=100%): /nvol_prd_db2_bku_01 1311004 MiB (27.78% inode=100%):
[1680772473] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680772487] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.414 second response time
[1680772533] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.097 second response time
[1680772552] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680772606] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772617] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680772671] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680772736] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680772762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1984 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6181 MiB (67.79% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4577 MiB (89.58% inode=100%): /nvol_prd_db2_bku_01 1304392 MiB (27.64% inode=100%):
[1680772769] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772812] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772814] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.408 second response time
[1680772833] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680772842] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680772876] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680772879] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680772898] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680772907] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680772969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 10 Hours, 57 Minutes
[1680772972] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773029] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773036] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773040] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680773062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1984 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4577 MiB (89.58% inode=100%): /nvol_prd_db2_bku_01 1299531 MiB (27.54% inode=100%):
[1680773094] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.378 second response time
[1680773162] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 1 Hours, 40 Minutes
[1680773241] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773306] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1984 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4577 MiB (89.58% inode=100%): /nvol_prd_db2_bku_01 1294326 MiB (27.43% inode=100%):
[1680773402] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.248 second response time
[1680773459] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773465] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773466] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773499] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773521] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773523] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680773530] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773564] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773574] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773586] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773588] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680773628] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773639] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773648] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773649] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773651] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1984 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4577 MiB (89.58% inode=100%): /nvol_prd_db2_bku_01 1289133 MiB (27.32% inode=100%):
[1680773693] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773703] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773713] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773714] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773715] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.433 second response time
[1680773757] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773768] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680773777] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773821] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680773832] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680773842] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680773869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 11 Hours, 12 Minutes
[1680773886] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680773892] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680773907] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1984 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4577 MiB (89.58% inode=100%): /nvol_prd_db2_bku_01 1283691 MiB (27.20% inode=100%):
[1680773971] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680773971] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680774018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.232 second response time
[1680774035] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680774036] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680774061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 1 Hours, 55 Minutes
[1680774100] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680774101] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680774103] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680774166] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680774168] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680774230] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680774262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1984 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4569 MiB (89.42% inode=100%): /nvol_prd_db2_bku_01 1277029 MiB (27.06% inode=100%):
[1680774318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.418 second response time
[1680774544] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680774562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1984 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1268239 MiB (26.87% inode=100%):
[1680774609] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680774619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.403 second response time
[1680774679] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680774705] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680774743] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680774769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 11 Hours, 27 Minutes
[1680774769] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680774833] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680774834] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680774862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1983 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1257302 MiB (26.64% inode=100%):
[1680774897] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680774918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.374 second response time
[1680774961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 2 Hours, 10 Minutes
[1680774962] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680775032] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680775095] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680775097] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680775138] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680775160] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680775162] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680775162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1983 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1243293 MiB (26.34% inode=100%):
[1680775219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.400 second response time
[1680775462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1193798 MiB (25.29% inode=100%):
[1680775518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.363 second response time
[1680775669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 11 Hours, 42 Minutes
[1680775762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1168472 MiB (24.76% inode=100%):
[1680775767] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680775818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.435 second response time
[1680775831] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680775861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 2 Hours, 25 Minutes
[1680776043] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680776062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1129311 MiB (23.93% inode=100%):
[1680776108] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680776119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.354 second response time
[1680776172] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680776236] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680776291] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680776356] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680776363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1119325 MiB (23.72% inode=100%):
[1680776418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.363 second response time
[1680776569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 11 Hours, 57 Minutes
[1680776641] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113670 MB (19.75% inode=100%):
[1680776662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1109392 MiB (23.51% inode=100%):
[1680776719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.395 second response time
[1680776762] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 2 Hours, 40 Minutes
[1680776837] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680776902] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680776962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1100238 MiB (23.31% inode=100%):
[1680777018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.610 second response time
[1680777024] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777089] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777096] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777135] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777161] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777175] SERVICE ALERT: vswprdaotap13;serGEN vswprd G4 Home REACT Service;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Proxy Error - string 'OK' not found on 'https://10.12.6.229:12443/healthz' - 559 bytes in 0.009 second response time
[1680777200] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777208] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777235] SERVICE ALERT: vswprdaotap13;serGEN vswprd G4 Home REACT Service;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 440 bytes in 0.012 second response time
[1680777262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4574 MiB (89.51% inode=100%): /nvol_prd_db2_bku_01 1159344 MiB (24.56% inode=100%):
[1680777273] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777295] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.377 second response time
[1680777359] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777382] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777447] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 12 Hours, 12 Minutes
[1680777500] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1951 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1150771 MiB (24.38% inode=100%):
[1680777565] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680777573] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.179 second response time
[1680777630] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680777638] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 2 Hours, 55 Minutes
[1680777680] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680777745] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680777862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1950 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1141893 MiB (24.19% inode=100%):
[1680777918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.412 second response time
[1680778054] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680778119] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680778162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1950 MiB (0.95% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1137415 MiB (24.10% inode=100%):
[1680778183] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680778219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.350 second response time
[1680778369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 12 Hours, 27 Minutes
[1680778412] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680778462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1918 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6183 MiB (67.81% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1135328 MiB (24.06% inode=100%):
[1680778477] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680778518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.478 second response time
[1680778561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 3 Hours, 10 Minutes
[1680778710] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680778762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1918 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1132879 MiB (24.00% inode=100%):
[1680778766] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680778775] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680778819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.433 second response time
[1680778831] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680779062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1918 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1129973 MiB (23.94% inode=100%):
[1680779118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.448 second response time
[1680779203] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680779267] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680779269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 12 Hours, 42 Minutes
[1680779362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1918 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1124617 MiB (23.83% inode=100%):
[1680779419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.384 second response time
[1680779434] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680779462] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 3 Hours, 25 Minutes
[1680779499] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680779662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1918 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1121192 MiB (23.76% inode=100%):
[1680779718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.528 second response time
[1680779742] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680779796] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680779807] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680779861] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680779962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1918 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1117653 MiB (23.68% inode=100%):
[1680780019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.344 second response time
[1680780066] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680780131] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680780169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 12 Hours, 57 Minutes
[1680780262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1918 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1114769 MiB (23.62% inode=100%):
[1680780318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.453 second response time
[1680780361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 3 Hours, 40 Minutes
[1680780471] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680780536] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680780562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1917 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1114364 MiB (23.61% inode=100%):
[1680780603] SERVICE DOWNTIME ALERT: jmplx_prd02_external_checks;serGEN vswprd mardb replication checks via jmplx_prd02;STARTED; Service has entered a period of scheduled downtime
[1680780618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.340 second response time
[1680780640] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680780705] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680780862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1917 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1113741 MiB (23.60% inode=100%):
[1680780919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.409 second response time
[1680781069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 13 Hours, 12 Minutes
[1680781132] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680781162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1917 MiB (0.93% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 1112922 MiB (23.58% inode=100%):
[1680781197] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680781218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.412 second response time
[1680781261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 3 Hours, 55 Minutes
[1680781462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1885 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4563 MiB (89.29% inode=100%): /nvol_prd_db2_bku_01 1111603 MiB (23.55% inode=100%):
[1680781519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.360 second response time
[1680781762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1885 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 1390042 MiB (29.45% inode=100%):
[1680781818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.464 second response time
[1680781969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 13 Hours, 27 Minutes
[1680782041] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113658 MB (19.75% inode=100%):
[1680782062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1885 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 1387372 MiB (29.40% inode=100%):
[1680782119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.390 second response time
[1680782161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 4 Hours, 10 Minutes
[1680782363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1885 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 1384900 MiB (29.34% inode=100%):
[1680782418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.486 second response time
[1680782662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1885 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.44% inode=100%): /nvol_prd_db2_bku_01 1382749 MiB (29.30% inode=100%):
[1680782719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.505 second response time
[1680782870] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 13 Hours, 42 Minutes
[1680782941] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113660 MB (19.75% inode=100%):
[1680782962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1885 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.44% inode=100%): /nvol_prd_db2_bku_01 1380275 MiB (29.25% inode=100%):
[1680783018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.485 second response time
[1680783062] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 4 Hours, 25 Minutes
[1680783262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1884 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.44% inode=100%): /nvol_prd_db2_bku_01 1376815 MiB (29.17% inode=100%):
[1680783319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.394 second response time
[1680783412] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680783472] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.070 second response time
[1680783526] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680783562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1884 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.44% inode=100%): /nvol_prd_db2_bku_01 1372768 MiB (29.09% inode=100%):
[1680783591] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680783618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.408 second response time
[1680783769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 13 Hours, 57 Minutes
[1680783841] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113668 MB (19.75% inode=100%):
[1680783862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1852 MiB (0.90% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 1371737 MiB (29.07% inode=100%):
[1680783918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.439 second response time
[1680783961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 4 Hours, 40 Minutes
[1680784162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1852 MiB (0.90% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 1366520 MiB (28.96% inode=100%):
[1680784219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.485 second response time
[1680784462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1820 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6182 MiB (67.80% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 1362629 MiB (28.87% inode=100%):
[1680784519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.600 second response time
[1680784574] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680784638] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680784669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 14 Hours, 12 Minutes
[1680784726] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680784741] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113660 MB (19.75% inode=100%):
[1680784762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1820 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4570 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2766424 MiB (58.62% inode=100%):
[1680784790] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680784818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.302 second response time
[1680784861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 4 Hours, 55 Minutes
[1680784966] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680785026] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578260 bytes in 0.097 second response time
[1680785062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1820 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2762801 MiB (58.55% inode=100%):
[1680785073] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680785074] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680785118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.251 second response time
[1680785134] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578259 bytes in 0.124 second response time
[1680785134] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578222 bytes in 0.164 second response time
[1680785362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1820 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2759031 MiB (58.47% inode=100%):
[1680785419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.457 second response time
[1680785569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 14 Hours, 27 Minutes
[1680785641] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113670 MB (19.75% inode=100%):
[1680785662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1819 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2750854 MiB (58.29% inode=100%):
[1680785700] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680785718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.406 second response time
[1680785762] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 5 Hours, 10 Minutes
[1680785765] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680785962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1819 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2742358 MiB (58.11% inode=100%):
[1680786019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.533 second response time
[1680786262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1819 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2734804 MiB (57.95% inode=100%):
[1680786318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.340 second response time
[1680786469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 14 Hours, 42 Minutes
[1680786562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1819 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2726804 MiB (57.78% inode=100%):
[1680786618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.440 second response time
[1680786661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 5 Hours, 25 Minutes
[1680786862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1819 MiB (0.88% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2718356 MiB (57.60% inode=100%):
[1680786919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.443 second response time
[1680787137] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680787162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1787 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6173 MiB (67.70% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2710119 MiB (57.43% inode=100%):
[1680787197] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680787218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.424 second response time
[1680787257] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680787369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 14 Hours, 57 Minutes
[1680787441] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113668 MB (19.75% inode=100%):
[1680787462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1787 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2704016 MiB (57.30% inode=100%):
[1680787519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.680 second response time
[1680787561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 5 Hours, 40 Minutes
[1680787762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1787 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2703015 MiB (57.28% inode=100%):
[1680787819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.542 second response time
[1680788063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1787 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2697492 MiB (57.16% inode=100%):
[1680788118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.451 second response time
[1680788269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 15 Hours, 12 Minutes
[1680788341] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113670 MB (19.75% inode=100%):
[1680788362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1787 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4566 MiB (89.36% inode=100%): /nvol_prd_db2_bku_01 2697492 MiB (57.16% inode=100%):
[1680788418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.398 second response time
[1680788461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 5 Hours, 55 Minutes
[1680788662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1786 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4564 MiB (89.31% inode=100%): /nvol_prd_db2_bku_01 2697491 MiB (57.16% inode=100%):
[1680788719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.425 second response time
[1680788962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1786 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4558 MiB (89.19% inode=100%): /nvol_prd_db2_bku_01 2697491 MiB (57.16% inode=100%):
[1680789028] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680789169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 15 Hours, 27 Minutes
[1680789262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1786 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.19% inode=100%): /nvol_prd_db2_bku_01 2697491 MiB (57.16% inode=100%):
[1680789318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.452 second response time
[1680789362] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 6 Hours, 10 Minutes
[1680789562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1786 MiB (0.87% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.19% inode=100%): /nvol_prd_db2_bku_01 2697490 MiB (57.16% inode=100%):
[1680789619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.476 second response time
[1680789928] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680790069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 15 Hours, 42 Minutes
[1680790162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1754 MiB (0.85% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.19% inode=100%): /nvol_prd_db2_bku_01 2697489 MiB (57.16% inode=100%):
[1680790219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.419 second response time
[1680790261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 6 Hours, 25 Minutes
[1680790462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1754 MiB (0.85% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.19% inode=100%): /nvol_prd_db2_bku_01 2697489 MiB (57.16% inode=100%):
[1680790528] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680790762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1754 MiB (0.85% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.97% inode=100%): /nvol_prd_db2_bku_01 2697489 MiB (57.16% inode=100%):
[1680790969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 15 Hours, 57 Minutes
[1680791161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 6 Hours, 40 Minutes
[1680791362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1753 MiB (0.85% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.97% inode=100%): /nvol_prd_db2_bku_01 2697489 MiB (57.16% inode=100%):
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: jmplx_prd02_external_checks;serGEN vswprd mardb replication checks via jmplx_prd02;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791401] SERVICE DOWNTIME ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791402] SERVICE DOWNTIME ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791402] SERVICE DOWNTIME ALERT: www-status;serGEN vswprd F5 Sessions;STOPPED; Service has exited from a period of scheduled downtime
[1680791402] SERVICE DOWNTIME ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680791420] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.298 second response time
[1680791719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.423 second response time
[1680791869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 16 Hours, 12 Minutes
[1680791962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1753 MiB (0.85% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.96% inode=100%): /nvol_prd_db2_bku_01 2697489 MiB (57.16% inode=100%):
[1680792019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.333 second response time
[1680792061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 6 Hours, 55 Minutes
[1680792262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1753 MiB (0.85% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697489 MiB (57.16% inode=100%):
[1680792320] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.440 second response time
[1680792619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.381 second response time
[1680792769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 16 Hours, 27 Minutes
[1680792862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1721 MiB (0.84% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697489 MiB (57.16% inode=100%):
[1680792920] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.368 second response time
[1680792962] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 7 Hours, 10 Minutes
[1680793219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.294 second response time
[1680793462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1721 MiB (0.84% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2696976 MiB (57.15% inode=100%):
[1680793519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.492 second response time
[1680793669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 16 Hours, 42 Minutes
[1680793763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1721 MiB (0.84% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.89% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2696976 MiB (57.15% inode=100%):
[1680793828] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680793861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 7 Hours, 25 Minutes
[1680794062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1720 MiB (0.84% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2696976 MiB (57.15% inode=100%):
[1680794351] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680794362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1720 MiB (0.84% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2696981 MiB (57.15% inode=100%):
[1680794411] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.057 second response time
[1680794420] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.372 second response time
[1680794569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 16 Hours, 57 Minutes
[1680794719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.289 second response time
[1680794761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 7 Hours, 40 Minutes
[1680794962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1720 MiB (0.84% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680795028] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680795195] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680795259] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680795469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 17 Hours, 12 Minutes
[1680795562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1720 MiB (0.84% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680795619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.415 second response time
[1680795662] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 7 Hours, 55 Minutes
[1680795862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1688 MiB (0.82% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6180 MiB (67.77% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4589 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680795918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.354 second response time
[1680796162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1698 MiB (0.82% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4589 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680796218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.421 second response time
[1680796369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 17 Hours, 27 Minutes
[1680796519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.630 second response time
[1680796561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 8 Hours, 10 Minutes
[1680796713] SERVICE ALERT: jboss_heap;serGEN vswprd jboss-heap-authservices;WARNING;SOFT;1;Warning - Warning on svpjbshc14.prd.allegiantair.com
[1680796828] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680796884] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680796884] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680796888] SERVICE ALERT: vswprdwpxwa02;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680796894] SERVICE ALERT: jboss_heap;serGEN vswprd jboss-heap-authservices;WARNING;SOFT;2;Warning - Warning on svpjbshc13.prd.allegiantair.com
[1680796944] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578222 bytes in 0.214 second response time
[1680796944] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578319 bytes in 0.112 second response time
[1680796948] SERVICE ALERT: vswprdwpxwa02;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578259 bytes in 0.099 second response time
[1680797062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1697 MiB (0.82% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4589 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680797074] SERVICE ALERT: jboss_heap;serGEN vswprd jboss-heap-authservices;OK;SOFT;3;OK -
[1680797269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 17 Hours, 42 Minutes
[1680797362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1697 MiB (0.82% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6190 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680797419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.420 second response time
[1680797461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 8 Hours, 25 Minutes
[1680797662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1697 MiB (0.82% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680797719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.391 second response time
[1680798028] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680798169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 17 Hours, 57 Minutes
[1680798318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.351 second response time
[1680798361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 8 Hours, 40 Minutes
[1680798562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1665 MiB (0.81% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680798619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.568 second response time
[1680798862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1665 MiB (0.81% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680798918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.556 second response time
[1680799070] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 18 Hours, 12 Minutes
[1680799162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1665 MiB (0.81% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680799219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.529 second response time
[1680799262] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 8 Hours, 55 Minutes
[1680799463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1664 MiB (0.81% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680799528] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680799652] SERVICE ALERT: svpkayws11;serGEN vswprd Kayak BLI to LAS;WARNING;SOFT;1;(No output on stdout) stderr: Traceback (most recent call last):
[1680799712] SERVICE ALERT: svpkayws11;serGEN vswprd Kayak BLI to LAS;OK;SOFT;2;OK - Kayak service healthy
[1680799955] SERVICE ALERT: jboss_heap;serGEN vswprd jboss-heap-authservices;WARNING;SOFT;1;Warning - Warning on svpjbshc13.prd.allegiantair.com
[1680799969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 18 Hours, 27 Minutes
[1680800062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1632 MiB (0.79% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680800136] SERVICE ALERT: jboss_heap;serGEN vswprd jboss-heap-authservices;OK;SOFT;2;OK -
[1680800161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 9 Hours, 10 Minutes
[1680800362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1632 MiB (0.79% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680800418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.460 second response time
[1680800512] SERVICE DOWNTIME ALERT: external_checks;serGEN vswprd Check g4-prod.airtrfx.com/en/;STARTED; Service has entered a period of scheduled downtime
[1680800512] SERVICE DOWNTIME ALERT: external_checks;serGEN vswprd Check prodshows.wpengine.com;STARTED; Service has entered a period of scheduled downtime
[1680800512] HOST DOWNTIME ALERT: vswtsesxi01;STARTED; Host has entered a period of scheduled downtime
[1680800512] SERVICE DOWNTIME ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;STARTED; Service has entered a period of scheduled downtime
[1680800512] SERVICE DOWNTIME ALERT: external_checks;serGEN vswprd Check www.allegiantair.jobs;STARTED; Service has entered a period of scheduled downtime
[1680800719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.439 second response time
[1680800869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 18 Hours, 42 Minutes
[1680800962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1600 MiB (0.78% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6181 MiB (67.79% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680801018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.556 second response time
[1680801061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 9 Hours, 25 Minutes
[1680801262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1600 MiB (0.78% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680801319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.418 second response time
[1680801618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.495 second response time
[1680801769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 18 Hours, 57 Minutes
[1680801862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1600 MiB (0.78% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680801919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.501 second response time
[1680801962] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 9 Hours, 40 Minutes
[1680802162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1600 MiB (0.78% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680802218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.510 second response time
[1680802377] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;1;DISK WARNING - free space: /var/lib/mongo 246059 MB (20.00% inode=100%):
[1680802463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1599 MiB (0.78% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680802519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.477 second response time
[1680802557] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;2;DISK WARNING - free space: /var/lib/mongo 246059 MB (20.00% inode=100%):
[1680802669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 19 Hours, 12 Minutes
[1680802737] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;3;DISK WARNING - free space: /var/lib/mongo 246060 MB (20.00% inode=100%):
[1680802818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.421 second response time
[1680802861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 9 Hours, 55 Minutes
[1680802917] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;4;DISK WARNING - free space: /var/lib/mongo 246058 MB (20.00% inode=100%):
[1680803062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1599 MiB (0.78% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680803097] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;OK;HARD;5;DISK OK - free space: /var/lib/mongo 246185 MB (20.01% inode=100%):
[1680803118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.445 second response time
[1680803362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1599 MiB (0.78% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680803419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.453 second response time
[1680803569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 19 Hours, 27 Minutes
[1680803662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1567 MiB (0.76% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680803718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.420 second response time
[1680803761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 10 Hours, 10 Minutes
[1680803962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1567 MiB (0.76% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680803997] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;1;DISK WARNING - free space: /var/lib/mongo 246056 MB (20.00% inode=100%):
[1680804019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.601 second response time
[1680804177] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;OK;SOFT;2;DISK OK - free space: /var/lib/mongo 246088 MB (20.01% inode=100%):
[1680804318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.510 second response time
[1680804469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 19 Hours, 42 Minutes
[1680804619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.477 second response time
[1680804661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 10 Hours, 25 Minutes
[1680804918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.395 second response time
[1680805162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1566 MiB (0.76% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680805219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.598 second response time
[1680805290] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680805350] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.068 second response time
[1680805369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 19 Hours, 57 Minutes
[1680805462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1566 MiB (0.76% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680805518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.377 second response time
[1680805562] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 10 Hours, 40 Minutes
[1680805762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1566 MiB (0.76% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.74% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680805819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.384 second response time
[1680806062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1566 MiB (0.76% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4587 MiB (89.77% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680806118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.368 second response time
[1680806157] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;1;DISK WARNING - free space: /var/lib/mongo 246054 MB (20.00% inode=100%):
[1680806270] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 20 Hours, 12 Minutes
[1680806337] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;OK;SOFT;2;DISK OK - free space: /var/lib/mongo 246255 MB (20.02% inode=100%):
[1680806362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1534 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6163 MiB (67.59% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680806419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.878 second response time
[1680806461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 10 Hours, 55 Minutes
[1680806517] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;1;DISK WARNING - free space: /var/lib/mongo 246054 MB (20.00% inode=100%):
[1680806662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1534 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680806697] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;OK;SOFT;2;DISK OK - free space: /var/lib/mongo 246085 MB (20.01% inode=100%):
[1680806718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.487 second response time
[1680806877] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;1;DISK WARNING - free space: /var/lib/mongo 246053 MB (20.00% inode=100%):
[1680806962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1534 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680807019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.400 second response time
[1680807057] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;2;DISK WARNING - free space: /var/lib/mongo 246053 MB (20.00% inode=100%):
[1680807169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 20 Hours, 27 Minutes
[1680807237] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;OK;SOFT;3;DISK OK - free space: /var/lib/mongo 246086 MB (20.01% inode=100%):
[1680807262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1534 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680807318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.464 second response time
[1680807361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 11 Hours, 10 Minutes
[1680807562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1534 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680807619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.381 second response time
[1680807862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1533 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680807918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.535 second response time
[1680808069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 20 Hours, 42 Minutes
[1680808162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1533 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680808219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.360 second response time
[1680808262] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 11 Hours, 25 Minutes
[1680808317] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;1;DISK WARNING - free space: /var/lib/mongo 245892 MB (19.99% inode=100%):
[1680808463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1533 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680808497] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;2;DISK WARNING - free space: /var/lib/mongo 245850 MB (19.99% inode=100%):
[1680808518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.526 second response time
[1680808677] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;3;DISK WARNING - free space: /var/lib/mongo 245882 MB (19.99% inode=100%):
[1680808762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1533 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680808818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.221 second response time
[1680808857] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;SOFT;4;DISK WARNING - free space: /var/lib/mongo 245850 MB (19.99% inode=100%):
[1680808969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 20 Hours, 57 Minutes
[1680809037] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;HARD;5;DISK WARNING - free space: /var/lib/mongo 246060 MB (20.00% inode=100%):
[1680809062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1501 MiB (0.73% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680809119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.218 second response time
[1680809161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 11 Hours, 40 Minutes
[1680809362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1501 MiB (0.73% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680809418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.528 second response time
[1680809662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1501 MiB (0.73% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680809719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.465 second response time
[1680809869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 21 Hours, 12 Minutes
[1680809937] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;HARD;5;DISK WARNING - free space: /var/lib/mongo 245823 MB (19.98% inode=100%):
[1680809962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1501 MiB (0.73% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680810018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.460 second response time
[1680810061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 11 Hours, 55 Minutes
[1680810262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1510 MiB (0.73% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680810318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.186 second response time
[1680810619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.413 second response time
[1680810769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 21 Hours, 27 Minutes
[1680810862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1510 MiB (0.73% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680810918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.396 second response time
[1680810961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 12 Hours, 10 Minutes
[1680811162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1510 MiB (0.73% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680811219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.526 second response time
[1680811462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1478 MiB (0.72% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6179 MiB (67.76% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4574 MiB (89.52% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680811518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.416 second response time
[1680811669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 21 Hours, 42 Minutes
[1680811762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1478 MiB (0.72% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4582 MiB (89.67% inode=100%): /nvol_prd_db2_bku_01 2697003 MiB (57.15% inode=100%):
[1680811818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.325 second response time
[1680811862] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 12 Hours, 25 Minutes
[1680812119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.353 second response time
[1680812418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.432 second response time
[1680812569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 21 Hours, 57 Minutes
[1680812637] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;HARD;5;DISK WARNING - free space: /var/lib/mongo 245858 MB (19.99% inode=100%):
[1680812662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1478 MiB (0.72% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4582 MiB (89.67% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680812719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.527 second response time
[1680812761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 12 Hours, 40 Minutes
[1680812962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1478 MiB (0.72% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4582 MiB (89.67% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680813018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.550 second response time
[1680813262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1477 MiB (0.72% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.83% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4582 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680813319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.345 second response time
[1680813469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 22 Hours, 12 Minutes
[1680813537] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;WARNING;HARD;5;DISK WARNING - free space: /var/lib/mongo 245736 MB (19.98% inode=100%):
[1680813562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1477 MiB (0.72% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4582 MiB (89.66% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680813618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.263 second response time
[1680813661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 12 Hours, 55 Minutes
[1680813862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1445 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680813918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.420 second response time
[1680813954] SERVICE ALERT: vswprdmonor01;serGEN vswprd Volume - Var-Lib-Mongo;OK;HARD;5;DISK OK - free space: /var/lib/mongo 450547 MB (31.40% inode=100%):
[1680814163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680814219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.377 second response time
[1680814369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 22 Hours, 27 Minutes
[1680814518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.531 second response time
[1680814562] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 13 Hours, 10 Minutes
[1680814762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680814819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.482 second response time
[1680815062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680815118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.233 second response time
[1680815269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 22 Hours, 42 Minutes
[1680815418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.405 second response time
[1680815461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 13 Hours, 25 Minutes
[1680815719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.560 second response time
[1680815962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1412 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680816018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.429 second response time
[1680816169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 22 Hours, 57 Minutes
[1680816229] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680816262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1412 MiB (0.68% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680816262] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul charter-customer;CRITICAL;SOFT;1;Critical - All nodes return passing. Critical on totalnodes
[1680816289] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.060 second response time
[1680816301] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816303] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.436 second response time
[1680816322] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul charter-customer;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816323] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 13 Hours, 40 Minutes
[1680816362] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816363] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816365] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816382] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul charter-customer;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816384] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816384] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816385] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft-mx;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816387] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816391] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816391] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul station;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816393] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul charter-customer;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816394] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816395] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul airport;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816416] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft-mx;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816417] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816417] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816422] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816423] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816425] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816438] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul airport;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816438] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul station;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816440] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816444] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816445] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816446] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft-mx;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816447] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816451] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816452] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul station;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816453] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul charter-customer;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816454] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816456] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul airport;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816476] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft-mx;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816477] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816477] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816479] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816479] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul charter-customer;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816486] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816499] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul airport;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816499] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul station;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816500] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft-mx;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816500] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816505] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816506] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft-mx;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816508] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816512] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816512] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul station;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816514] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul charter-customer;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816515] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816516] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul airport;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816520] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul airport;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816521] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul station;UNKNOWN;SOFT;1;Unknown - length of json array 0
[1680816536] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft-mx;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816538] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-order-importer-web;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816538] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-boot-admin;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816539] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816539] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul charter-customer;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816559] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul airport;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816559] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul station;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816560] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-cloud-config;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816560] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft-mx;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1380 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6180 MiB (67.78% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680816581] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul airport;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816581] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul station;UNKNOWN;SOFT;2;Unknown - length of json array 0
[1680816600] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul charter-customer;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816600] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul acars-gw;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul aircraft;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul aircraft-mx;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul airport;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul cc-event-processor;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul charter-customer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul check for leader agreement;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul dmx-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-actual-times-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-base-management;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-disposition;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-disposition-orch;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-info;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-jms;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-record;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-adf-surrogate;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-aircraft-mx-event-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-aircraft-mx-surrogate;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-aircraft-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-airport-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-anomaly-finder-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-config;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-flight-ops-data-orchestrator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-flight-ops-data-orchestrator-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-flight-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-irop-manager;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-tail-assigner;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-ui-gateway;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-ui-gateway-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-flight-tax;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-loyalty-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-ops-seats;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-boot-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-cloud-config;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-accounting;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-aos;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-line;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-lmto;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-materials;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-mco;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-pi;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-portal;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-bulk-order-cancel;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-flight-order;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-flight-shop;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-flight-status;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-session-manager;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul order-payment-event-processor;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-customer-residency;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-discount;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-checkin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-status;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-analyzer-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-custcoms;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-event-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-exporter-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-ext;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-flight-alerts;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-history;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-importer-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-package-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-travelagent-commission;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul policy;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul profile-cc-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul profile-traveler-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul station;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul zipkin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul zone-boarding;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul acars-gw;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul aircraft;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul aircraft-mx;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul airport;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul cc-event-processor;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul charter-customer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul check for leader agreement;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul dmx-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-actual-times-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-base-management;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-disposition;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-disposition-orch;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-info;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-jms;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-record;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-adf-surrogate;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-aircraft-mx-event-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-aircraft-mx-surrogate;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-aircraft-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-airport-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-anomaly-finder-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-config;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-flight-ops-data-orchestrator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-flight-ops-data-orchestrator-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-flight-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-irop-manager;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-tail-assigner;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-ui-gateway;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-ui-gateway-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-flight-tax;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-loyalty-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-ops-seats;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-boot-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-cloud-config;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-accounting;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-aos;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-line;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-lmto;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-materials;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-mco;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-pi;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-portal;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-bulk-order-cancel;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-flight-order;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-flight-shop;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-flight-status;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-session-manager;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul order-payment-event-processor;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-customer-residency;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-discount;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-checkin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-status;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-analyzer-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-custcoms;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-event-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-exporter-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-ext;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-flight-alerts;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-history;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-importer-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-package-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-travelagent-commission;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul policy;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul profile-cc-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul profile-traveler-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul station;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul zipkin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul zone-boarding;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul acars-gw;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul aircraft;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul aircraft-mx;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul airport;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul cc-event-processor;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul charter-customer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul check for leader agreement;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul dmx-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-actual-times-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-base-management;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-disposition;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-disposition-orch;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-info;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-jms;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-record;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-adf-surrogate;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-aircraft-mx-event-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-aircraft-mx-surrogate;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-aircraft-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-airport-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-anomaly-finder-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-config;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-flight-ops-data-orchestrator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-flight-ops-data-orchestrator-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-flight-translator;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-irop-manager;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-tail-assigner;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-ui-gateway;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-ui-gateway-v2;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-flight-tax;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-loyalty-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-ops-seats;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-boot-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-cloud-config;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-accounting;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-admin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-aos;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-line;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-lmto;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-materials;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-mco;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-pi;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-portal;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-bulk-order-cancel;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-flight-order;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-flight-shop;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-flight-status;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-session-manager;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul order-payment-event-processor;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-customer-residency;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-discount;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-checkin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-status;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-analyzer-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-custcoms;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-event-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-exporter-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-ext;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-flight-alerts;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-history;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-importer-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-package-consumer;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-travelagent-commission;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul policy;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul profile-cc-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul profile-traveler-web;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul station;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul zipkin;STARTED; Service has entered a period of scheduled downtime
[1680816610] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul zone-boarding;STARTED; Service has entered a period of scheduled downtime
[1680816618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.470 second response time
[1680816620] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft-mx;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816641] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul airport;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816641] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul station;UNKNOWN;HARD;3;Unknown - length of json array 0
[1680816805] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-order-importer-web;CRITICAL;HARD;3;Critical - ota-order-importer-web-64639b5e4af7deff3869b840a0faea1a critical:. Critical on totalnodes. Critical on badpercent
[1680816807] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft-mx;CRITICAL;HARD;3;Critical - aircraft-mx-vswprdspnws01 critical:. Critical on totalnodes. Critical on badpercent
[1680816808] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-boot-admin;CRITICAL;HARD;3;Critical - g4-spring-boot-admin-vswprdspnws01 critical:. Critical on totalnodes. Critical on badpercent
[1680816812] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft;CRITICAL;HARD;3;Critical - All nodes return passing. Critical on totalnodes
[1680816812] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul station;CRITICAL;HARD;3;Critical - station-vswprdspnws01 critical:. Critical on totalnodes. Critical on badpercent
[1680816814] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul charter-customer;CRITICAL;HARD;3;Critical - All nodes return passing. Critical on totalnodes
[1680816815] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-cloud-config;CRITICAL;HARD;3;Critical - All nodes return passing. Critical on totalnodes
[1680816816] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul airport;CRITICAL;HARD;3;Critical - All nodes return passing. Critical on totalnodes
[1680816837] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft-mx;CRITICAL;HARD;3;Critical - All nodes return passing. Critical on totalnodes
[1680816838] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-order-importer-web;CRITICAL;HARD;3;Critical - All nodes return passing. Critical on totalnodes
[1680816838] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-boot-admin;CRITICAL;HARD;3;Critical - All nodes return passing. Critical on totalnodes
[1680816858] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft;OK;HARD;3;OK - All nodes return passing
[1680816859] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul airport;OK;HARD;3;OK - All nodes return passing
[1680816859] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul station;CRITICAL;HARD;3;Critical - station-vswprdspnws02 critical:. Critical on badpercent
[1680816860] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-cloud-config;CRITICAL;HARD;3;Critical - g4-spring-cloud-config-vswprdspnws02 critical:. Critical on badpercent
[1680816862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1380 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4579 MiB (89.61% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680816870] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul aircraft-mx;OK;HARD;3;OK - All nodes return passing
[1680816878] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-boot-admin;OK;HARD;3;OK - All nodes return passing
[1680816885] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul charter-customer;OK;HARD;3;OK - All nodes return passing
[1680816888] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul station;OK;HARD;3;OK - All nodes return passing
[1680816895] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft;OK;HARD;3;OK - All nodes return passing
[1680816900] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft;OK;HARD;3;OK - All nodes return passing
[1680816900] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul charter-customer;OK;HARD;3;OK - All nodes return passing
[1680816919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.332 second response time
[1680816920] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul aircraft-mx;OK;HARD;3;OK - All nodes return passing
[1680816929] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-cloud-config;OK;HARD;3;OK - All nodes return passing
[1680816930] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-boot-admin;OK;HARD;3;OK - All nodes return passing
[1680816939] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul airport;OK;HARD;3;OK - All nodes return passing
[1680816946] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-order-importer-web;OK;HARD;3;OK - All nodes return passing
[1680816948] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul airport;OK;HARD;3;OK - All nodes return passing
[1680816966] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul flight-disposition;WARNING;SOFT;1;Warning - flight-disposition-0d42b7179989a25d688dd2b0b6876cfa critical:. Warning on badpercent
[1680816970] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-boot-admin;OK;HARD;3;OK - All nodes return passing
[1680816982] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul charter-customer;OK;HARD;3;OK - All nodes return passing
[1680816982] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-order-importer-web;OK;HARD;3;OK - All nodes return passing
[1680816985] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-cloud-config;OK;HARD;3;OK - All nodes return passing
[1680817009] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-order-importer-web;OK;HARD;3;OK - All nodes return passing
[1680817021] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul flight-disposition;OK;SOFT;2;OK - All nodes return passing
[1680817027] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul station;OK;HARD;3;OK - All nodes return passing
[1680817034] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul station;OK;HARD;3;OK - All nodes return passing
[1680817042] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul aircraft-mx;OK;HARD;3;OK - All nodes return passing
[1680817054] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-cloud-config;OK;HARD;3;OK - All nodes return passing
[1680817069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 23 Hours, 12 Minutes
[1680817218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.287 second response time
[1680817261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 13 Hours, 55 Minutes
[1680817462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1380 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4575 MiB (89.53% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680817518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.500 second response time
[1680817655] SERVICE ALERT: svpsrywb11;serGEN vswprd ais maintenance page hit count;WARNING;SOFT;1;WARN: 29 page hits in the last minute to maintenance.allegiantair.com
[1680817656] SERVICE ALERT: svpsrywb11;serGEN vswprd sorry page count;WARNING;SOFT;1;WARN: 29 page hits in the last minute to maintenance.allegiantair.com
[1680817716] SERVICE ALERT: svpsrywb11;serGEN vswprd ais maintenance page hit count;OK;SOFT;2;OK - 0 unique users on maintenance.allegiantair.com in the last minute
[1680817716] SERVICE ALERT: svpsrywb11;serGEN vswprd sorry page count;OK;SOFT;2;OK - 0 unique users on maintenance.allegiantair.com in the last minute
[1680817762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1380 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680817819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.452 second response time
[1680817969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 23 Hours, 27 Minutes
[1680818062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1380 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680818118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.432 second response time
[1680818162] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 14 Hours, 10 Minutes
[1680818362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1380 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680818419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.344 second response time
[1680818461] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680818526] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680818662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1379 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680818718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.463 second response time
[1680818869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 23 Hours, 42 Minutes
[1680818962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1347 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680818981] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;SOFT;1;Warning - Warning on ActiveMQ.DLQ:634
[1680819018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.351 second response time
[1680819041] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;OK;SOFT;2;OK - No QueueSize issues
[1680819061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 14 Hours, 25 Minutes
[1680819262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1347 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680819319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.511 second response time
[1680819618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.426 second response time
[1680819769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 106 Days, 23 Hours, 57 Minutes
[1680819862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1347 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680819919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.350 second response time
[1680819961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 14 Hours, 40 Minutes
[1680820002] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;SOFT;1;Warning - Warning on ActiveMQ.DLQ:738
[1680820062] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;SOFT;2;Warning - Warning on ActiveMQ.DLQ:738
[1680820123] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;CRITICAL;HARD;3;Critical - Critical on ActiveMQ.DLQ:1192
[1680820163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1347 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680820218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.438 second response time
[1680820423] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:590
[1680820496] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680820519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.557 second response time
[1680820556] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578296 bytes in 0.192 second response time
[1680820669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 12 Hours, 12 Minutes
[1680820762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1347 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680820818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.569 second response time
[1680820862] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 14 Hours, 55 Minutes
[1680821062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1347 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680821119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.447 second response time
[1680821362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1346 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680821418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.456 second response time
[1680821569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 27 Hours, 27 Minutes
[1680821662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1314 MiB (0.64% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6164 MiB (67.60% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.43% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680821719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.468 second response time
[1680821761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 15 Hours, 10 Minutes
[1680821924] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:733
[1680821962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1314 MiB (0.64% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680822018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.372 second response time
[1680822223] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:744
[1680822319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.549 second response time
[1680822470] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 42 Hours, 42 Minutes
[1680822524] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:756
[1680822562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1314 MiB (0.64% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680822618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.465 second response time
[1680822661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 15 Hours, 25 Minutes
[1680822823] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:754
[1680822862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1314 MiB (0.64% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680822918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.353 second response time
[1680823123] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:750
[1680823219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.461 second response time
[1680823364] SERVICE ALERT: vswprdg4pmx01;serGEN vswprd dmx g4pmx-pi proxy check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Proxy Error - 579 bytes in 0.010 second response time
[1680823369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 57 Hours, 57 Minutes
[1680823424] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:922
[1680823424] SERVICE ALERT: vswprdg4pmx01;serGEN vswprd dmx g4pmx-pi proxy check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 804 bytes in 0.037 second response time
[1680823462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1314 MiB (0.64% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680823463] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114084 MB (19.82% inode=100%):
[1680823518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.464 second response time
[1680823561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 15 Hours, 40 Minutes
[1680823723] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:874
[1680823752] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul dmx-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823752] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-history;STOPPED; Service has exited from a period of scheduled downtime
[1680823752] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul zone-boarding;STOPPED; Service has exited from a period of scheduled downtime
[1680823752] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-aircraft-mx-surrogate;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-analyzer-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-status;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-flight-status;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-exporter-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul zone-boarding;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul zipkin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul station;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul profile-traveler-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul profile-cc-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul policy;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-lmto;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-flight-order;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-disposition-orch;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-disposition;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul aircraft;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul acars-gw;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-base-management;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul cc-event-processor;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul airport;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul aircraft-mx;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-actual-times-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul charter-customer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul dmx-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul check for leader agreement;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-bulk-order-cancel;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-tail-assigner;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-aircraft-mx-surrogate;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-info;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-aircraft-mx-event-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-jms;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-adf-surrogate;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul flight-record;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-irop-manager;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-anomaly-finder-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-airport-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-aircraft-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-flight-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-config;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-flight-ops-data-orchestrator-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-flight-ops-data-orchestrator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-portal;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-cloud-config;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-ui-gateway-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul fmm-ui-gateway;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-spring-boot-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-flight-tax;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-ops-seats;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4-loyalty-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-pi;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-aos;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-accounting;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-mco;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-line;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-materials;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul g4mx-lmto;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-line;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul cc-event-processor;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-ext;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-customer-residency;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul order-payment-event-processor;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-flight-shop;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-session-manager;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul navitaire-flight-status;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-exporter-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-checkin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-discount;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-event-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-ops-seats;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-status;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-custcoms;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-analyzer-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-ui-gateway;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-ui-gateway-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul airport;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul policy;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-flight-ops-data-orchestrator-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-history;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-flight-alerts;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-travelagent-commission;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-importer-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul ota-order-package-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul aircraft-mx;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-adf-surrogate;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul station;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul profile-traveler-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul profile-cc-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-disposition-orch;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul aircraft;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul zipkin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul acars-gw;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap01;serGEN vswprd Consul zone-boarding;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-aos;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-airport-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-disposition;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul charter-customer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-base-management;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul check for leader agreement;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-actual-times-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul dmx-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-aircraft-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-jms;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-info;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-aircraft-mx-surrogate;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul flight-record;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-aircraft-mx-event-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-tail-assigner;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-config;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-anomaly-finder-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-irop-manager;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-flight-ops-data-orchestrator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul fmm-flight-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-accounting;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-flight-tax;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul profile-cc-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-cloud-config;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-loyalty-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4-spring-boot-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-travelagent-commission;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-loyalty-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul station;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-checkin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-portal;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-pi;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-info;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-materials;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-checkin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-tail-assigner;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-package-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-session-manager;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-discount;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-session-manager;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-travelagent-commission;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul aircraft;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-event-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul aircraft-mx;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-custcoms;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-custcoms;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-accounting;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-ext;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-lmto;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-boot-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-flight-ops-data-orchestrator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-aircraft-mx-event-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-actual-times-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul check for leader agreement;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-bulk-order-cancel;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-customer-residency;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-flight-shop;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-line;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-base-management;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul charter-customer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-flight-status;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-aos;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-flight-order;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-package-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-ops-seats;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-analyzer-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul profile-traveler-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-disposition-orch;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-history;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-anomaly-finder-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-flight-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul g4mx-mco;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-ui-gateway-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-materials;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul cc-event-processor;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-flight-shop;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-pi;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul navitaire-bulk-order-cancel;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul order-payment-event-processor;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-event-consumer;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-ui-gateway;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul acars-gw;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-portal;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-admin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-flight-tax;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4mx-mco;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-adf-surrogate;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-flight-alerts;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-flight-alerts;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-irop-manager;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul navitaire-flight-order;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-aircraft-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-config;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-importer-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-airport-translator;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-customer-residency;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-discount;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul g4-spring-cloud-config;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-exporter-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-ext;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul fmm-flight-ops-data-orchestrator-v2;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-record;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul ota-order-importer-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul airport;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul policy;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-status;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-jms;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul ota-order-web;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap03;serGEN vswprd Consul flight-disposition;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul zipkin;STOPPED; Service has exited from a period of scheduled downtime
[1680823753] SERVICE DOWNTIME ALERT: vswprdconap02;serGEN vswprd Consul order-payment-event-processor;STOPPED; Service has exited from a period of scheduled downtime
[1680823762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1314 MiB (0.64% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680823819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.496 second response time
[1680824024] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:606
[1680824062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1282 MiB (0.62% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680824118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.409 second response time
[1680824269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 1 Hours, 12 Minutes
[1680824362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1281 MiB (0.62% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680824363] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114102 MB (19.82% inode=100%):
[1680824419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.425 second response time
[1680824462] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 15 Hours, 55 Minutes
[1680824661] SERVICE ALERT: vswprdg4pmx02;serGEN vswprd dmx g4pmx-pi proxy check;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 502 Proxy Error - 579 bytes in 0.009 second response time
[1680824662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1281 MiB (0.62% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680824721] SERVICE ALERT: vswprdg4pmx02;serGEN vswprd dmx g4pmx-pi proxy check;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 804 bytes in 0.030 second response time
[1680824924] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:646
[1680824962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1281 MiB (0.62% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680825019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.487 second response time
[1680825038] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680825102] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680825169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 1 Hours, 27 Minutes
[1680825223] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:782
[1680825262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1281 MiB (0.62% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680825263] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114087 MB (19.82% inode=100%):
[1680825318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.334 second response time
[1680825361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 16 Hours, 10 Minutes
[1680825524] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:784
[1680825563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1281 MiB (0.62% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680825619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.524 second response time
[1680825823] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:792
[1680825862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1281 MiB (0.62% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680825918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.561 second response time
[1680826069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 1 Hours, 42 Minutes
[1680826124] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:774
[1680826162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1249 MiB (0.61% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6180 MiB (67.77% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680826163] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114102 MB (19.82% inode=100%):
[1680826219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.530 second response time
[1680826261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 16 Hours, 25 Minutes
[1680826423] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:778
[1680826462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1249 MiB (0.61% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680826518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.395 second response time
[1680826723] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:746
[1680826762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1249 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680826818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.323 second response time
[1680826969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 1 Hours, 57 Minutes
[1680827024] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:778
[1680827062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1216 MiB (0.59% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680827119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.457 second response time
[1680827161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 16 Hours, 40 Minutes
[1680827168] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680827228] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.076 second response time
[1680827248] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827291] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827308] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827323] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:750
[1680827351] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1216 MiB (0.59% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680827368] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827411] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.336 second response time
[1680827457] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827517] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;OK;SOFT;2;OK - All nodes return passing
[1680827577] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827624] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:739
[1680827637] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1216 MiB (0.59% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680827698] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680827719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.432 second response time
[1680827869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 2 Hours, 12 Minutes
[1680827923] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:879
[1680827962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1216 MiB (0.59% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.45% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680827963] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114085 MB (19.82% inode=100%):
[1680827968] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680827998] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680828011] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680828018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.341 second response time
[1680828062] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 16 Hours, 55 Minutes
[1680828223] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:779
[1680828262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1216 MiB (0.59% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680828268] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828298] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828312] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.203 second response time
[1680828329] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;OK;SOFT;2;OK - All nodes return passing
[1680828358] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical:. Critical on badpercent
[1680828372] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828389] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828418] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828432] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828449] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680828509] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680828524] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:811
[1680828562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1184 MiB (0.57% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680828619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.483 second response time
[1680828718] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680828732] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680828769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 2 Hours, 27 Minutes
[1680828823] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:771
[1680828862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1184 MiB (0.57% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6188 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680828863] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114102 MB (19.82% inode=100%):
[1680828928] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680828961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 17 Hours, 10 Minutes
[1680829123] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:793
[1680829162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1184 MiB (0.57% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4568 MiB (89.39% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680829218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.257 second response time
[1680829424] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:741
[1680829528] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680829669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 2 Hours, 42 Minutes
[1680829723] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:737
[1680829762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1183 MiB (0.57% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4568 MiB (89.39% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680829861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 17 Hours, 25 Minutes
[1680830024] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:749
[1680830118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.419 second response time
[1680830323] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:735
[1680830362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1183 MiB (0.57% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4568 MiB (89.39% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680830419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.489 second response time
[1680830518] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830532] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 2 Hours, 57 Minutes
[1680830579] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830592] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830609] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830624] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:849
[1680830639] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830653] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1183 MiB (0.57% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4568 MiB (89.39% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680830664] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114092 MB (19.82% inode=100%):
[1680830670] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830728] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680830730] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;HARD;3;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680830762] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 17 Hours, 40 Minutes
[1680830923] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:755
[1680830939] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680830953] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680830963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1151 MiB (0.56% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6178 MiB (67.76% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4568 MiB (89.39% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680831030] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680831223] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:759
[1680831262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1151 MiB (0.56% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4567 MiB (89.39% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680831319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.455 second response time
[1680831469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 3 Hours, 12 Minutes
[1680831524] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:771
[1680831563] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114093 MB (19.82% inode=100%):
[1680831628] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680831661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 17 Hours, 55 Minutes
[1680831823] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:611
[1680831862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1151 MiB (0.56% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.85% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.31% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680831918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.323 second response time
[1680832124] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:753
[1680832162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1151 MiB (0.56% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4567 MiB (89.38% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680832228] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680832369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 3 Hours, 27 Minutes
[1680832423] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:875
[1680832462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1150 MiB (0.56% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.31% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680832463] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114099 MB (19.82% inode=100%):
[1680832553] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680832561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 18 Hours, 10 Minutes
[1680832617] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680832723] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:937
[1680832762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1150 MiB (0.56% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.31% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680832819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.341 second response time
[1680833024] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:617
[1680833062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1118 MiB (0.54% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680833118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.635 second response time
[1680833269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 3 Hours, 42 Minutes
[1680833323] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:619
[1680833363] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114102 MB (19.82% inode=100%):
[1680833428] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680833461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 18 Hours, 25 Minutes
[1680833624] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:601
[1680833662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1118 MiB (0.54% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680833719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.552 second response time
[1680833923] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:597
[1680833962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1118 MiB (0.54% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680834018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.517 second response time
[1680834169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 3 Hours, 57 Minutes
[1680834223] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:731
[1680834319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.539 second response time
[1680834362] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 18 Hours, 40 Minutes
[1680834524] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:717
[1680834562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1118 MiB (0.54% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680834618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.554 second response time
[1680834823] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:844
[1680834862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1118 MiB (0.54% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680834919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.387 second response time
[1680835024] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680835069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 4 Hours, 12 Minutes
[1680835088] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680835124] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:732
[1680835162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1085 MiB (0.53% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6163 MiB (67.59% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4563 MiB (89.30% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680835218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.456 second response time
[1680835261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 18 Hours, 55 Minutes
[1680835423] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:756
[1680835462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 733 MiB (0.35% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6161 MiB (67.57% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4597 MiB (89.96% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680835528] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680835723] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:616
[1680835762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 733 MiB (0.35% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680835819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.670 second response time
[1680835969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 4 Hours, 27 Minutes
[1680836024] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:726
[1680836119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.561 second response time
[1680836161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 19 Hours, 10 Minutes
[1680836323] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:658
[1680836362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 701 MiB (0.34% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680836418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.584 second response time
[1680836624] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:716
[1680836663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 701 MiB (0.34% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680836719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.504 second response time
[1680836869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 4 Hours, 42 Minutes
[1680836923] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:704
[1680836987] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680837018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.582 second response time
[1680837052] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680837062] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 19 Hours, 25 Minutes
[1680837223] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:786
[1680837262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 700 MiB (0.34% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680837328] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680837524] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:774
[1680837562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 700 MiB (0.34% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680837619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.508 second response time
[1680837770] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 4 Hours, 57 Minutes
[1680837823] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:758
[1680837918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.586 second response time
[1680837961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 19 Hours, 40 Minutes
[1680838107] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680838123] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:692
[1680838162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 668 MiB (0.32% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680838167] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.059 second response time
[1680838219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.591 second response time
[1680838369] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680838424] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:674
[1680838434] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680838462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 668 MiB (0.32% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680838518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.530 second response time
[1680838669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 5 Hours, 12 Minutes
[1680838723] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:604
[1680838819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.517 second response time
[1680838861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 19 Hours, 55 Minutes
[1680839024] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:600
[1680839062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 668 MiB (0.32% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680839119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.552 second response time
[1680839323] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:678
[1680839362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 668 MiB (0.32% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680839418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.515 second response time
[1680839569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 5 Hours, 27 Minutes
[1680839623] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:718
[1680839662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 636 MiB (0.31% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680839719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.456 second response time
[1680839761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 20 Hours, 10 Minutes
[1680839924] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:658
[1680839962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 635 MiB (0.31% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680840018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.409 second response time
[1680840224] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:720
[1680840262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 635 MiB (0.31% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680840318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.555 second response time
[1680840469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 5 Hours, 42 Minutes
[1680840525] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:716
[1680840562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 603 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6178 MiB (67.75% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680840619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.423 second response time
[1680840659] SERVICE ALERT: VSWPRDADSDC03;serGEN vswprd ldap certificate port 636;CRITICAL;SOFT;1;CRITICAL - Cannot make SSL connection.
[1680840662] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 20 Hours, 25 Minutes
[1680840719] SERVICE ALERT: VSWPRDADSDC03;serGEN vswprd ldap certificate port 636;OK;SOFT;2;SSL OK - Certificate 'vswprdadsdc03.allegiantair.com' will expire in 85 days on 2023-07-01 12:42 -0700/PDT.
[1680840813] SERVICE ALERT: svpimodb11;serGEN vswprd Inflight Number of Orders to be Processed;WARNING;SOFT;1;WARNING: count is 324
[1680840824] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:720
[1680840862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 603 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4594 MiB (89.91% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680840872] SERVICE ALERT: svpimodb12;serGEN vswprd Inflight Number of Orders to be Processed;WARNING;SOFT;1;WARNING: count is 370
[1680840873] SERVICE ALERT: svpimodb11;serGEN vswprd Inflight Number of Orders to be Processed;WARNING;SOFT;2;WARNING: count is 370
[1680840898] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680840918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.464 second response time
[1680840932] SERVICE ALERT: svpimodb12;serGEN vswprd Inflight Number of Orders to be Processed;OK;SOFT;2;OK: count is 270
[1680840933] SERVICE ALERT: svpimodb11;serGEN vswprd Inflight Number of Orders to be Processed;OK;HARD;3;OK: count is 270
[1680840963] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680841016] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680841081] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680841124] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:652
[1680841162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 603 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6186 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4594 MiB (89.91% inode=100%): /nvol_prd_db2_bku_01 2697002 MiB (57.15% inode=100%):
[1680841219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.711 second response time
[1680841369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 5 Hours, 57 Minutes
[1680841425] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:654
[1680841462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 603 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4594 MiB (89.91% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680841518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.364 second response time
[1680841561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 20 Hours, 40 Minutes
[1680841724] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:720
[1680841762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 603 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6184 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4594 MiB (89.91% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680841819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.445 second response time
[1680841967] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680842025] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:660
[1680842031] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680842062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 603 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4594 MiB (89.91% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680842118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.470 second response time
[1680842230] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680842269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 6 Hours, 12 Minutes
[1680842294] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680842324] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:656
[1680842362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 603 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4594 MiB (89.90% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680842419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.388 second response time
[1680842461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 20 Hours, 55 Minutes
[1680842624] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:642
[1680842662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 602 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6183 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680842718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.390 second response time
[1680842828] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680842892] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680842962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 602 MiB (0.29% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680843019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.387 second response time
[1680843169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 6 Hours, 27 Minutes
[1680843224] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:692
[1680843262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 570 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680843318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.499 second response time
[1680843362] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 21 Hours, 10 Minutes
[1680843491] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680843525] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:642
[1680843545] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680843555] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680843562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 570 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6183 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680843609] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680843619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.450 second response time
[1680843731] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680843795] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680843824] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:692
[1680843862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 570 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6185 MiB (67.84% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680843918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.470 second response time
[1680844069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 6 Hours, 42 Minutes
[1680844102] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680844124] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:626
[1680844167] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680844218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.351 second response time
[1680844232] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680844261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 21 Hours, 25 Minutes
[1680844297] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680844361] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680844425] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:632
[1680844425] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680844462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 570 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6183 MiB (67.82% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680844465] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680844519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.391 second response time
[1680844529] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680844556] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680844620] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680844626] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680844686] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched "200" - 578282 bytes in 0.254 second response time
[1680844724] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:634
[1680844739] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680844753] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680844756] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680844762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 570 MiB (0.27% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.83% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680844789] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680844800] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680844813] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;2;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680844818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.340 second response time
[1680844826] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1680844830] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;CRITICAL;SOFT;1;Critical - ota-flight-market-provider critical: ota-flight-market-provider critical:. Critical on badpercent
[1680844853] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680844860] SERVICE ALERT: vswprdconap01;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680844874] SERVICE ALERT: vswprdconap02;serGEN vswprd Consul ota-flight-market-provider;OK;HARD;3;OK - All nodes return passing
[1680844886] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;OK;HARD;3;HTTP OK: Status line output matched "200" - 578319 bytes in 0.109 second response time
[1680844891] SERVICE ALERT: vswprdconap03;serGEN vswprd Consul ota-flight-market-provider;OK;SOFT;2;OK - All nodes return passing
[1680844918] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680844969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 6 Hours, 57 Minutes
[1680845119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.400 second response time
[1680845161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 21 Hours, 40 Minutes
[1680845324] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:632
[1680845362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 538 MiB (0.26% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680845418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.433 second response time
[1680845493] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680845558] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680845662] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 537 MiB (0.26% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4591 MiB (89.85% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680845718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.368 second response time
[1680845822] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680845869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 7 Hours, 12 Minutes
[1680845881] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680845887] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680845895] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680845925] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:640
[1680845941] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113669 MB (19.75% inode=100%):
[1680845945] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680845960] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680846009] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.415 second response time
[1680846024] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680846061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 21 Hours, 55 Minutes
[1680846074] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680846143] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846208] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680846222] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846224] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:674
[1680846262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 537 MiB (0.26% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680846287] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680846525] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:670
[1680846562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 537 MiB (0.26% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680846573] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.497 second response time
[1680846629] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846637] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680846652] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846693] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680846716] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680846758] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 7 Hours, 27 Minutes
[1680846781] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680846822] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680846824] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:668
[1680846859] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680846887] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680846918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.457 second response time
[1680846923] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680846962] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 22 Hours, 10 Minutes
[1680847002] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847063] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847066] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680847124] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:660
[1680847127] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680847152] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 537 MiB (0.26% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2697001 MiB (57.15% inode=100%):
[1680847192] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680847216] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680847219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.448 second response time
[1680847220] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847281] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680847285] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680847425] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:658
[1680847462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 537 MiB (0.26% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2694795 MiB (57.11% inode=100%):
[1680847518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.311 second response time
[1680847669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 7 Hours, 42 Minutes
[1680847685] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847685] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847724] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:628
[1680847731] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847749] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680847750] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680847762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 515 MiB (0.25% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2689090 MiB (56.98% inode=100%):
[1680847796] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680847818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.507 second response time
[1680847819] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847860] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680847861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 22 Hours, 25 Minutes
[1680847884] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680847939] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680847949] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680847950] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848003] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680848014] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680848025] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:624
[1680848051] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848062] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 515 MiB (0.25% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2683292 MiB (56.86% inode=100%):
[1680848079] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680848116] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680848119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.389 second response time
[1680848185] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848249] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680848324] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:620
[1680848362] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 514 MiB (0.25% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2677483 MiB (56.74% inode=100%):
[1680848396] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.379 second response time
[1680848461] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680848488] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848525] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680848552] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680848569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 7 Hours, 57 Minutes
[1680848624] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:618
[1680848663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 514 MiB (0.25% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2671748 MiB (56.62% inode=100%):
[1680848718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.235 second response time
[1680848761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 22 Hours, 40 Minutes
[1680848765] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848829] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680848829] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680848916] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848925] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:628
[1680848962] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 514 MiB (0.25% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2666142 MiB (56.50% inode=100%):
[1680848981] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680848984] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680848995] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680849002] SERVICE DOWNTIME ALERT: www-status;serGEN vswprd F5 Sessions;STARTED; Service has entered a period of scheduled downtime
[1680849019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.519 second response time
[1680849046] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680849049] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849060] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680849106] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.062 second response time
[1680849114] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849124] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680849153] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849178] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849218] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849224] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:618
[1680849256] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849262] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 514 MiB (0.25% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 2660245 MiB (56.37% inode=100%):
[1680849318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.344 second response time
[1680849321] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849373] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849429] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680849438] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 8 Hours, 12 Minutes
[1680849525] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:612
[1680849562] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 514 MiB (0.25% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 2654443 MiB (56.25% inode=100%):
[1680849567] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849583] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.539 second response time
[1680849631] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849647] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 22 Hours, 55 Minutes
[1680849680] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849745] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849802] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849813] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680849824] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:610
[1680849862] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 482 MiB (0.23% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6187 MiB (67.86% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 2648640 MiB (56.13% inode=100%):
[1680849867] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849878] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680849918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.396 second response time
[1680850125] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:608
[1680850162] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 482 MiB (0.23% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 2606445 MiB (55.23% inode=100%):
[1680850192] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850214] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.341 second response time
[1680850256] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680850279] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680850321] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 107 Days, 8 Hours, 27 Minutes
[1680850385] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680850409] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850424] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:610
[1680850450] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850462] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 482 MiB (0.23% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 2552096 MiB (54.08% inode=100%):
[1680850474] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680850514] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680850518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.330 second response time
[1680850562] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 372 Days, 23 Hours, 10 Minutes
[1680850633] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850644] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850693] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680850697] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680850708] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680850724] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq DLQ QueueSize;WARNING;HARD;3;Warning - Warning on ActiveMQ.DLQ:604
[1680850758] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680850762] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680850762] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 481 MiB (0.23% inode=98%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4585 MiB (89.73% inode=100%): /nvol_prd_db2_bku_01 2498283 MiB (52.94% inode=100%):
[1680850772] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result