[1680505219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.518 second response time
[1680505356] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680505421] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680505426] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680505463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1436 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4589 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1836842 MiB (38.92% inode=100%):
[1680505486] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680505518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.680 second response time
[1680505551] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680505617] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680505669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 8 Hours, 42 Minutes"
[1680505763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1436 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4593 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 1784923 MiB (37.82% inode=100%):
[1680505819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.286 second response time
[1680505852] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680505861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 368 Days, 23 Hours, 25 Minutes"
[1680505902] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680505917] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680505920] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680505967] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680505982] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680505985] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506004] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1436 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.86% inode=100%): /nvol_prd_db2_bku_01 1728472 MiB (36.63% inode=100%):
[1680506064] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506069] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.350 second response time
[1680506130] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506134] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506196] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506200] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506223] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506261] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680506288] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680506288] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506327] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680506335] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1446 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.86% inode=100%): /nvol_prd_db2_bku_01 1669479 MiB (35.38% inode=100%):
[1680506400] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680506418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.476 second response time
[1680506466] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680506569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 8 Hours, 57 Minutes"
[1680506650] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1446 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.86% inode=100%): /nvol_prd_db2_bku_01 1624984 MiB (34.43% inode=100%):
[1680506715] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.492 second response time
"[1680506761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 368 Days, 23 Hours, 40 Minutes"
[1680506833] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506865] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506898] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506930] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680506931] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506932] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680506963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1446 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.86% inode=100%): /nvol_prd_db2_bku_01 1595747 MiB (33.81% inode=100%):
[1680506986] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680506997] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680507000] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.403 second response time
[1680507051] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680507062] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680507065] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680507098] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507117] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680507163] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680507228] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680507233] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507237] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1446 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.86% inode=100%): /nvol_prd_db2_bku_01 1573331 MiB (33.34% inode=100%):
[1680507298] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680507303] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680507319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.408 second response time
[1680507364] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680507368] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507422] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680507433] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680507469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 9 Hours, 12 Minutes"
[1680507494] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507559] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680507563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1445 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.86% inode=100%): /nvol_prd_db2_bku_01 1552153 MiB (32.89% inode=100%):
[1680507564] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.498 second response time
[1680507624] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507629] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680507662] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 368 Days, 23 Hours, 55 Minutes"
[1680507668] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507670] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507689] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680507729] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680507734] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680507735] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680507755] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680507795] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680507799] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680507863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1445 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1533094 MiB (32.49% inode=100%):
[1680507919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.424 second response time
[1680508060] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680508133] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680508163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1445 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1514538 MiB (32.09% inode=100%):
[1680508195] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680508218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.394 second response time
[1680508260] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680508326] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680508327] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680508365] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680508369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 9 Hours, 27 Minutes"
[1680508393] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680508431] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680508438] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680508464] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1494588 MiB (31.67% inode=100%):
[1680508504] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680508519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.530 second response time
"[1680508561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 10 Hours, 10 Minutes"
[1680508569] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680508635] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680508763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4590 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1475003 MiB (31.25% inode=100%):
[1680508797] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680508818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.581 second response time
[1680508826] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680508863] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680508872] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680508897] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680508938] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13618 MiB (88.71% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.31% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1459292 MiB (30.92% inode=100%):
[1680509119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.528 second response time
[1680509180] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509231] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509231] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509245] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680509269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 9 Hours, 42 Minutes"
[1680509296] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509296] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509304] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509353] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509358] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4590 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1443905 MiB (30.60% inode=100%):
[1680509369] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.436 second response time
[1680509419] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509424] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680509461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 25 Hours, 25 Minutes"
[1680509529] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509539] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509562] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509594] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509604] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509628] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680509663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1413 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4590 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1422762 MiB (30.15% inode=100%):
[1680509670] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509692] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680509719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.522 second response time
[1680509735] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509735] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509800] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680509800] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509842] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509866] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680509866] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680509900] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509907] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680509910] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509931] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680509933] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680509961] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680509963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1412 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4589 MiB (89.82% inode=100%): /nvol_prd_db2_bku_01 1402678 MiB (29.72% inode=100%):
[1680509965] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680509972] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680509976] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680509993] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.079 second response time
[1680510018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.341 second response time
[1680510027] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510031] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680510093] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510158] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
"[1680510169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 9 Hours, 57 Minutes"
[1680510224] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680510260] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510261] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1412 MiB (0.68% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4589 MiB (89.81% inode=100%): /nvol_prd_db2_bku_01 1385923 MiB (29.37% inode=100%):
[1680510318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.452 second response time
[1680510325] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510326] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510336] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680510341] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680510361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 40 Hours, 40 Minutes"
[1680510384] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510406] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680510449] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510471] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680510500] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510504] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510529] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510536] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2813 MiB (1.37% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4589 MiB (89.81% inode=100%): /nvol_prd_db2_bku_01 1378649 MiB (29.21% inode=100%):
[1680510565] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510569] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510594] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510598] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510601] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680510619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.506 second response time
[1680510635] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510663] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510666] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680510690] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510691] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510700] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510755] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680510757] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680510821] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680510863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2801 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.31% inode=100%): /var/log/audit 4589 MiB (89.81% inode=100%): /nvol_prd_db2_bku_01 1372316 MiB (29.08% inode=100%):
[1680510918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.420 second response time
[1680510931] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510942] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510991] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680510997] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511007] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511028] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511057] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511062] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680511069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 10 Hours, 12 Minutes"
[1680511072] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680511076] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511093] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511122] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511142] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511159] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680511163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2801 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.31% inode=100%): /var/log/audit 4589 MiB (89.81% inode=100%): /nvol_prd_db2_bku_01 1365759 MiB (28.94% inode=100%):
[1680511187] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.469 second response time
[1680511253] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680511262] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 55 Hours, 55 Minutes"
[1680511368] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511372] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511433] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511437] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2801 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1359723 MiB (28.81% inode=100%):
[1680511464] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.364 second response time
[1680511529] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511558] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511571] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511596] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680511623] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511637] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511688] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511726] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511754] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2800 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1352721 MiB (28.66% inode=100%):
[1680511777] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511791] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511798] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511807] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.570 second response time
[1680511843] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511857] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680511859] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680511863] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680511872] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
"[1680511892] SERVICE ALERT: vswprdprxap01;serGEN vswprd RHEL Load;WARNING;SOFT;1;WARNING - load average: 3.32, 1.20, 0.51"
[1680511908] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680511924] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680511929] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680511938] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680511952] SERVICE ALERT: vswprdprxap01;serGEN vswprd RHEL Load;OK;SOFT;2;OK - load average: 1.79, 1.17, 0.55"
"[1680511969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 10 Hours, 27 Minutes"
[1680511977] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680512043] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680512063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2800 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1344969 MiB (28.50% inode=100%):
[1680512103] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680512118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.426 second response time
"[1680512161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 1 Hours, 10 Minutes"
[1680512162] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680512168] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680512228] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680512233] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680512266] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680512289] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680512315] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680512331] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680512355] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680512363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2800 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1337460 MiB (28.34% inode=100%):
[1680512380] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680512419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.379 second response time
[1680512445] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680512513] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680512659] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680512663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2800 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1329618 MiB (28.17% inode=100%):
[1680512718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.579 second response time
[1680512724] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680512869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 10 Hours, 42 Minutes"
[1680512963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2800 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1322103 MiB (28.01% inode=100%):
[1680512996] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.407 second response time
[1680513021] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513051] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680513061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 1 Hours, 25 Minutes"
[1680513062] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680513086] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513089] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513094] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513116] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513127] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680513138] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513155] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513160] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513204] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513225] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2800 MiB (1.36% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1314353 MiB (27.85% inode=100%):
[1680513290] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.452 second response time
[1680513401] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513435] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513466] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513500] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513520] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513564] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2768 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1307272 MiB (27.70% inode=100%):
[1680513569] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513586] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.483 second response time
[1680513634] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680513770] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 10 Hours, 57 Minutes"
[1680513822] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2768 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 1302373 MiB (27.60% inode=100%):
[1680513888] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680513908] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680513918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.489 second response time
"[1680513962] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 1 Hours, 40 Minutes"
[1680513974] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514039] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514081] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514098] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514103] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514105] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514131] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514146] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2768 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4582 MiB (89.67% inode=100%): /nvol_prd_db2_bku_01 1297345 MiB (27.49% inode=100%):
[1680514163] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514168] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514196] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680514219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.433 second response time
[1680514229] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514251] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514256] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514262] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680514294] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680514299] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514316] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680514321] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514360] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680514365] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514382] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680514463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2768 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4582 MiB (89.67% inode=100%): /nvol_prd_db2_bku_01 1292454 MiB (27.39% inode=100%):
[1680514518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.519 second response time
[1680514554] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514618] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514619] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514633] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514651] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680514669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 11 Hours, 12 Minutes"
[1680514683] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680514698] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514717] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514730] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514749] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680514763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2767 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.99% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1287266 MiB (27.28% inode=100%):
[1680514771] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514782] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514796] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680514811] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.449 second response time
[1680514836] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514847] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680514861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 1 Hours, 55 Minutes"
[1680514861] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680514876] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680514985] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514986] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680514987] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680515050] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680515051] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680515053] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680515063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2767 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1280523 MiB (27.13% inode=100%):
[1680515119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.475 second response time
[1680515266] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680515363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2767 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1271881 MiB (26.95% inode=100%):
[1680515418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.436 second response time
[1680515467] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680515512] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680515532] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680515569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 11 Hours, 27 Minutes"
[1680515571] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680515577] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680515594] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680515637] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680515643] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680515659] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680515663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2767 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1264013 MiB (26.78% inode=100%):
[1680515719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.449 second response time
"[1680515761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 2 Hours, 10 Minutes"
[1680515801] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680515867] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680515932] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680515963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2767 MiB (1.35% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6199 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1257353 MiB (26.64% inode=100%):
[1680516003] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680516018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.444 second response time
[1680516025] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680516068] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680516090] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680516156] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680516221] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680516263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2735 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1214789 MiB (25.74% inode=100%):
[1680516402] SERVICE ALERT: vswprdaisap11;serGEN vswprd AIS HTTPS Connect;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680516462] SERVICE ALERT: vswprdaisap11;serGEN vswprd AIS HTTPS Connect;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 3463 bytes in 0.044 second response time
"[1680516469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 11 Hours, 42 Minutes"
[1680516563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2735 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1179795 MiB (25.00% inode=100%):
[1680516586] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680516618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.477 second response time
[1680516633] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680516652] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
"[1680516661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 2 Hours, 25 Minutes"
[1680516698] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680516717] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680516733] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680516798] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680516838] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680516863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2735 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4578 MiB (89.60% inode=100%): /nvol_prd_db2_bku_01 1137511 MiB (24.10% inode=100%):
[1680516864] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680516903] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680516919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.470 second response time
[1680517023] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680517088] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680517097] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680517163] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680517163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2735 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1130675 MiB (23.96% inode=100%):
[1680517215] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680517218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.523 second response time
[1680517228] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680517281] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680517369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 11 Hours, 57 Minutes"
[1680517441] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113672 MB (19.75% inode=100%):
[1680517463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2734 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1121415 MiB (23.76% inode=100%):
[1680517469] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680517519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.535 second response time
[1680517534] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680517562] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 2 Hours, 40 Minutes"
[1680517600] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680517665] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680517744] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680517763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2734 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1111978 MiB (23.56% inode=100%):
"[1680517804] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578259 bytes in 0.085 second response time"
[1680517818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.375 second response time
[1680518063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2734 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1148828 MiB (24.34% inode=100%):
[1680518119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.415 second response time
[1680518247] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680518269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 12 Hours, 12 Minutes"
[1680518312] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680518363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2734 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1140857 MiB (24.17% inode=100%):
[1680518378] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680518418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.438 second response time
[1680518443] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680518461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 2 Hours, 55 Minutes"
[1680518663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2734 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1137517 MiB (24.10% inode=100%):
[1680518718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.409 second response time
[1680518953] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680518964] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2734 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1134837 MiB (24.05% inode=100%):
[1680519019] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680519019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.440 second response time
[1680519163] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680519169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 12 Hours, 27 Minutes"
[1680519229] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680519263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2734 MiB (1.33% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1129548 MiB (23.93% inode=100%):
[1680519318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.388 second response time
"[1680519361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 3 Hours, 10 Minutes"
[1680519563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2702 MiB (1.31% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1125773 MiB (23.85% inode=100%):
[1680519619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.528 second response time
[1680519863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2702 MiB (1.31% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1122409 MiB (23.78% inode=100%):
[1680519918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.396 second response time
"[1680520069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 12 Hours, 42 Minutes"
[1680520163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1118925 MiB (23.71% inode=100%):
[1680520219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.409 second response time
"[1680520262] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 3 Hours, 25 Minutes"
[1680520413] SERVICE ALERT: vswprdmardb02;serGEN vswprd DBA MySQL Replication Slave Status;WARNING;SOFT;1;WARNING: Slave is 177 seconds behind Master
[1680520455] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680520463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1117544 MiB (23.68% inode=100%):
[1680520473] SERVICE ALERT: vswprdmardb02;serGEN vswprd DBA MySQL Replication Slave Status;OK;SOFT;2;OK: Slave SQL running: Yes Slave IO running: Yes / master: 10.12.8.71 / slave is 89 seconds behind master
[1680520518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.581 second response time
[1680520520] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680520533] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680520585] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680520585] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680520599] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680520640] SERVICE ALERT: jmplx_prd02_external_checks;serGEN vswprd mardb replication checks via jmplx_prd02;WARNING;SOFT;1;WARNING: Slave is 143 seconds behind Master
[1680520650] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680520651] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680520701] SERVICE ALERT: jmplx_prd02_external_checks;serGEN vswprd mardb replication checks via jmplx_prd02;WARNING;SOFT;2;WARNING: Slave is 204 seconds behind Master
[1680520762] SERVICE ALERT: jmplx_prd02_external_checks;serGEN vswprd mardb replication checks via jmplx_prd02;OK;HARD;3;OK: Slave SQL running: Yes Slave IO running: Yes / master: 10.12.8.71 / slave is 0 seconds behind master
[1680520763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1116912 MiB (23.67% inode=100%):
[1680520766] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680520819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.453 second response time
[1680520831] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680520872] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680520933] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.069 second response time
[1680520965] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680520969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 12 Hours, 57 Minutes"
[1680521016] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680521030] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680521063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1115958 MiB (23.65% inode=100%):
[1680521082] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680521118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.211 second response time
"[1680521161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 3 Hours, 40 Minutes"
[1680521363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1114716 MiB (23.62% inode=100%):
[1680521402] SERVICE DOWNTIME ALERT: jmplx_prd02_external_checks;serGEN vswprd mardb replication checks via jmplx_prd02;STARTED; Service has entered a period of scheduled downtime
[1680521418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.440 second response time
[1680521663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1113181 MiB (23.59% inode=100%):
[1680521695] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680521719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.441 second response time
[1680521760] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680521869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 13 Hours, 12 Minutes"
[1680521963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 1110958 MiB (23.54% inode=100%):
[1680522018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.439 second response time
"[1680522061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 3 Hours, 55 Minutes"
[1680522263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2669 MiB (1.30% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1389223 MiB (29.44% inode=100%):
[1680522319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.413 second response time
[1680522563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2637 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10159 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1386590 MiB (29.38% inode=100%):
[1680522618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.418 second response time
"[1680522769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 13 Hours, 27 Minutes"
[1680522863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1383714 MiB (29.32% inode=100%):
[1680522919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.415 second response time
"[1680522961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 4 Hours, 10 Minutes"
[1680522964] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;SOFT;1;WARNING: Found 20 lines (limit=1/0): 03 Apr 2023 04:55:54 FATAL Inside handleError
[1680523144] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;SOFT;2;WARNING: Found 24 lines (limit=1/0): 03 Apr 2023 04:59:04 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680523163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1380324 MiB (29.25% inode=100%):
[1680523218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.348 second response time
[1680523324] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;SOFT;3;WARNING: Found 24 lines (limit=1/0): 03 Apr 2023 05:02:04 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680523463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1377060 MiB (29.18% inode=100%):
[1680523504] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;SOFT;4;WARNING: Found 24 lines (limit=1/0): 03 Apr 2023 05:05:04 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680523518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.471 second response time
"[1680523669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 13 Hours, 42 Minutes"
[1680523684] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 24 lines (limit=1/0): 03 Apr 2023 05:08:04 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680523763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1375748 MiB (29.15% inode=100%):
[1680523819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.469 second response time
"[1680523862] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 4 Hours, 25 Minutes"
[1680524063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1371339 MiB (29.06% inode=100%):
[1680524118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.472 second response time
[1680524364] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1367141 MiB (28.97% inode=100%):
[1680524419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.561 second response time
"[1680524569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 13 Hours, 57 Minutes"
[1680524585] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 120 lines (limit=1/0): 03 Apr 2023 05:23:04 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680524663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1363327 MiB (28.89% inode=100%):
[1680524718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.479 second response time
"[1680524761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 4 Hours, 40 Minutes"
[1680524963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1359860 MiB (28.81% inode=100%):
[1680525019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.585 second response time
[1680525263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2636 MiB (1.28% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 1356974 MiB (28.75% inode=100%):
[1680525318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.275 second response time
"[1680525469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 14 Hours, 12 Minutes"
[1680525484] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 120 lines (limit=1/0): 03 Apr 2023 05:38:05 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680525563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2604 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 2534447 MiB (53.71% inode=100%):
[1680525619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.434 second response time
"[1680525661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 4 Hours, 55 Minutes"
[1680525833] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680525863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 2526101 MiB (53.53% inode=100%):
"[1680525893] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578311 bytes in 0.308 second response time"
[1680525918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.450 second response time
[1680526163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 2517954 MiB (53.36% inode=100%):
[1680526219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.402 second response time
"[1680526369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 14 Hours, 27 Minutes"
[1680526384] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 120 lines (limit=1/0): 03 Apr 2023 05:53:05 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680526463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2509671 MiB (53.18% inode=100%):
[1680526518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.564 second response time
"[1680526562] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 5 Hours, 10 Minutes"
[1680526763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2501476 MiB (53.01% inode=100%):
[1680526819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.502 second response time
[1680527063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2493706 MiB (52.84% inode=100%):
[1680527118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.411 second response time
"[1680527270] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 14 Hours, 42 Minutes"
[1680527284] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 120 lines (limit=1/0): 03 Apr 2023 06:08:05 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680527363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2485904 MiB (52.68% inode=100%):
[1680527419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.508 second response time
"[1680527461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 5 Hours, 25 Minutes"
[1680527663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2483729 MiB (52.63% inode=100%):
[1680527718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.506 second response time
[1680527963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2603 MiB (1.27% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2481301 MiB (52.58% inode=100%):
[1680528018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.251 second response time
"[1680528169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 14 Hours, 57 Minutes"
[1680528185] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 120 lines (limit=1/0): 03 Apr 2023 06:23:04 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680528263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2571 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2477726 MiB (52.50% inode=100%):
[1680528319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.458 second response time
"[1680528361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 5 Hours, 40 Minutes"
[1680528563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2570 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2477726 MiB (52.50% inode=100%):
[1680528618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.411 second response time
[1680528863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2570 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4557 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2477726 MiB (52.50% inode=100%):
[1680528919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.456 second response time
"[1680529069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 15 Hours, 12 Minutes"
[1680529084] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 182 lines (limit=1/0): 03 Apr 2023 06:38:00 INFO  could not resetConnection inside initSocketHandler swallow and proceedError could not open connection to Server at 172.16.4.77:9601
[1680529163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2570 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4556 MiB (89.17% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680529218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.439 second response time
"[1680529261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 5 Hours, 55 Minutes"
[1680529463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2570 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680529519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.617 second response time
[1680529818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.475 second response time
"[1680529969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 15 Hours, 27 Minutes"
"[1680529984] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;WARNING;HARD;5;WARNING: Found 133 lines (limit=1/0): 03 Apr 2023 06:50:51 FATAL loggging off exceptionError reading from this.in, this.in = null"
[1680530063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2570 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680530119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.524 second response time
"[1680530162] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 6 Hours, 10 Minutes"
[1680530363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2570 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680530418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.466 second response time
[1680530663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2570 MiB (1.25% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680530719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.566 second response time
[1680530785] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680530846] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578236 bytes in 0.315 second response time"
"[1680530869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 15 Hours, 42 Minutes"
[1680530884] SERVICE ALERT: vswprdaisap02;serGEN vswprd Baglink Connection;OK;HARD;5;OK: Found 0 lines (limit=1/0): No matches found.
[1680530963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2537 MiB (1.23% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680531018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.464 second response time
"[1680531061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 6 Hours, 25 Minutes"
[1680531263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2537 MiB (1.23% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680531319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.466 second response time
[1680531619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.748 second response time
[1680531663] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680531723] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578260 bytes in 0.330 second response time"
"[1680531769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 15 Hours, 57 Minutes"
[1680531812] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680531863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2537 MiB (1.23% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680531872] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.058 second response time
[1680531918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.451 second response time
"[1680531961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 6 Hours, 40 Minutes"
[1680532200] SERVICE DOWNTIME ALERT: www-status;serGEN vswprd F5 Sessions;STOPPED; Service has exited from a period of scheduled downtime
[1680532200] SERVICE DOWNTIME ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532200] SERVICE DOWNTIME ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532200] SERVICE DOWNTIME ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532200] SERVICE DOWNTIME ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532200] SERVICE DOWNTIME ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: jmplx_prd02_external_checks;serGEN vswprd mardb replication checks via jmplx_prd02;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532201] SERVICE DOWNTIME ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;STOPPED; Service has exited from a period of scheduled downtime
[1680532218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.470 second response time
[1680532463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2537 MiB (1.23% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.98% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680532519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.478 second response time
"[1680532669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 16 Hours, 12 Minutes"
[1680532818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.552 second response time
"[1680532862] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 6 Hours, 55 Minutes"
[1680533063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2537 MiB (1.23% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680533119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.397 second response time
[1680533418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.412 second response time
"[1680533569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 16 Hours, 27 Minutes"
[1680533719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.387 second response time
"[1680533761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 7 Hours, 10 Minutes"
[1680533963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2504 MiB (1.22% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477725 MiB (52.50% inode=100%):
[1680534018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.648 second response time
[1680534263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2504 MiB (1.22% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477214 MiB (52.49% inode=100%):
[1680534319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.471 second response time
"[1680534469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 16 Hours, 42 Minutes"
[1680534618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.476 second response time
"[1680534661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 7 Hours, 25 Minutes"
[1680534920] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.964 second response time
[1680535090] SERVICE ALERT: vswprdjboss18;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 200 OK - 1765 bytes in 25.690 second response time
[1680535118] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;WARNING;SOFT;1;Warning - Warning on CustComsAsyncQueue:198
[1680535145] SERVICE ALERT: vswprdjboss17;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 200 OK - 1765 bytes in 11.939 second response time
[1680535164] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2504 MiB (1.22% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477221 MiB (52.49% inode=100%):
[1680535169] SERVICE ALERT: vswprdjboss20;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680535178] SERVICE ALERT: vswprdjboss18;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1680535178] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;SOFT;2;Critical - Critical on CustComsAsyncQueue:283
[1680535208] SERVICE ALERT: vswprdjboss17;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 1765 bytes in 3.367 second response time
[1680535228] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1680535235] SERVICE ALERT: vswprdjboss19;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 200 OK - 1765 bytes in 18.402 second response time
[1680535239] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:377
[1680535242] SERVICE ALERT: vswprdjboss20;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;CRITICAL;SOFT;2;CRITICAL - Cannot make SSL connection.
[1680535249] SERVICE ALERT: vswprdjboss18;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 200 OK - 1765 bytes in 11.324 second response time
[1680535295] SERVICE ALERT: vswprdjboss19;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 1765 bytes in 0.279 second response time
[1680535302] SERVICE ALERT: vswprdjboss20;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 1765 bytes in 0.282 second response time
"[1680535369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 16 Hours, 57 Minutes"
[1680535436] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:580
[1680535463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2504 MiB (1.22% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2477220 MiB (52.49% inode=100%):
[1680535519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.490 second response time
[1680535550] SERVICE ALERT: vswprdjboss18;serGEN vswprd Cheetah service at api.eccmp.com cheetah api;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 1765 bytes in 0.282 second response time
"[1680535561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 7 Hours, 40 Minutes"
[1680535737] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:922
[1680535763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2504 MiB (1.22% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4554 MiB (89.13% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680535818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.559 second response time
[1680536036] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;WARNING;HARD;3;Warning - Warning on CustComsAsyncQueue:162
[1680536119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.569 second response time
"[1680536269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 17 Hours, 12 Minutes"
[1680536325] SERVICE ALERT: vswprdjboss17;JbossSaGEN prod custcoms_web;CRITICAL;SOFT;1;connect to address 10.12.6.172 and port 8280: Connection refused
[1680536336] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:536
[1680536385] SERVICE ALERT: vswprdjboss17;JbossSaGEN prod custcoms_web;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 354 bytes in 0.007 second response time
[1680536418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.517 second response time
"[1680536462] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 7 Hours, 55 Minutes"
[1680536637] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:732
[1680536663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2471 MiB (1.20% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6179 MiB (67.77% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4586 MiB (89.75% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680536719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.471 second response time
[1680536936] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:1142
[1680536963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2471 MiB (1.20% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.97% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680537018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.454 second response time
"[1680537169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 17 Hours, 27 Minutes"
[1680537237] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:1414
[1680537263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2439 MiB (1.19% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.97% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680537319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.501 second response time
"[1680537361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 8 Hours, 10 Minutes"
[1680537536] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:1679
[1680537563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2439 MiB (1.19% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.96% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680537618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.440 second response time
[1680537836] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:2376
[1680537840] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Rsyslog work directory;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 10 seconds.
[1680537863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2439 MiB (1.19% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4597 MiB (89.96% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680537900] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Rsyslog work directory;OK;SOFT;2;OK
[1680537919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.390 second response time
"[1680538069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 17 Hours, 42 Minutes"
[1680538137] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:2726
[1680538163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2439 MiB (1.19% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.96% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680538218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.479 second response time
"[1680538261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 8 Hours, 25 Minutes"
[1680538436] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:3013
[1680538463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2439 MiB (1.19% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.96% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680538519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.431 second response time
[1680538737] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:3503
[1680538818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.440 second response time
"[1680538969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 17 Hours, 57 Minutes"
[1680539036] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:3848
[1680539063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2439 MiB (1.19% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4597 MiB (89.96% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
"[1680539161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 8 Hours, 40 Minutes"
[1680539337] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:4255
[1680539363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2406 MiB (1.17% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680539418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.542 second response time
[1680539636] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:4801
[1680539663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.18% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680539719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.426 second response time
[1680539815] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680539816] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680539869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 18 Hours, 12 Minutes"
"[1680539876] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578296 bytes in 0.354 second response time"
"[1680539876] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578282 bytes in 0.378 second response time"
[1680539937] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;CRITICAL;HARD;3;Critical - Critical on CustComsAsyncQueue:4607
[1680539963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.18% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6198 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680540018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.676 second response time
"[1680540061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 8 Hours, 55 Minutes"
[1680540236] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;WARNING;HARD;3;Warning - Warning on CustComsAsyncQueue:145
[1680540264] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.18% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4593 MiB (89.89% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680540319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.475 second response time
[1680540536] SERVICE ALERT: vswprdamqap11;serGEN vswprd amq QueueSize;OK;HARD;3;OK - No QueueSize issues
[1680540563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.18% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680540618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.502 second response time
"[1680540769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 18 Hours, 27 Minutes"
[1680540863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.18% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4593 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680540919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.802 second response time
"[1680540961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 9 Hours, 10 Minutes"
[1680541163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.18% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680541218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.458 second response time
[1680541463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.17% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680541519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.632 second response time
"[1680541670] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 18 Hours, 42 Minutes"
[1680541763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2416 MiB (1.17% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680541819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.518 second response time
"[1680541861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 9 Hours, 25 Minutes"
[1680542063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2384 MiB (1.16% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4593 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680542118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.653 second response time
[1680542363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2383 MiB (1.16% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4594 MiB (89.90% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680542419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.507 second response time
"[1680542569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 18 Hours, 57 Minutes"
[1680542663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2383 MiB (1.16% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4594 MiB (89.90% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680542718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.589 second response time
[1680542750] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680542761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 9 Hours, 40 Minutes"
[1680542810] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.060 second response time
[1680543019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.571 second response time
[1680543319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.569 second response time
"[1680543469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 19 Hours, 12 Minutes"
[1680543618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.507 second response time
"[1680543661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 9 Hours, 55 Minutes"
[1680543863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2383 MiB (1.16% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4591 MiB (89.84% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680543919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.546 second response time
[1680544163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2383 MiB (1.16% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4591 MiB (89.84% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680544218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.631 second response time
"[1680544369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 19 Hours, 27 Minutes"
[1680544519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.501 second response time
"[1680544561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 10 Hours, 10 Minutes"
[1680544763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2361 MiB (1.15% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.97% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4591 MiB (89.84% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680544818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.508 second response time
[1680545063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2360 MiB (1.15% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4591 MiB (89.84% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680545119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.571 second response time
"[1680545269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 19 Hours, 42 Minutes"
[1680545363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2360 MiB (1.15% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4591 MiB (89.84% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680545418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.504 second response time
"[1680545461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 10 Hours, 25 Minutes"
[1680545719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.695 second response time
[1680546019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.552 second response time
"[1680546169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 19 Hours, 57 Minutes"
[1680546318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.691 second response time
"[1680546362] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 10 Hours, 40 Minutes"
[1680546563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2360 MiB (1.15% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4591 MiB (89.84% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680546619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.464 second response time
[1680546918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.459 second response time
"[1680547069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 20 Hours, 12 Minutes"
[1680547163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2328 MiB (1.13% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.30% inode=100%): /var/log/audit 4591 MiB (89.84% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680547219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.529 second response time
"[1680547261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 10 Hours, 55 Minutes"
[1680547463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2328 MiB (1.13% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.77% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680547518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.461 second response time
[1680547716] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680547763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2327 MiB (1.13% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.77% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
"[1680547771] SERVICE ALERT: vswprdaisap14;Aerodata soap service;CRITICAL;HARD;1;Network failure, cannot contact website"
"[1680547776] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578258 bytes in 0.095 second response time"
[1680547819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.457 second response time
[1680547831] SERVICE ALERT: vswprdaisap14;Aerodata soap service;OK;HARD;1;HTTP/1.1 200 OK
"[1680547969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 20 Hours, 27 Minutes"
[1680548063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2327 MiB (1.13% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.77% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680548118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.456 second response time
"[1680548161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 11 Hours, 10 Minutes"
[1680548419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.510 second response time
[1680548663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2327 MiB (1.13% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680548718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.518 second response time
"[1680548869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 20 Hours, 42 Minutes"
[1680548963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2327 MiB (1.13% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680549019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.523 second response time
"[1680549061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 11 Hours, 25 Minutes"
[1680549318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.449 second response time
[1680549563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2327 MiB (1.13% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680549619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.532 second response time
[1680549674] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680549734] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578222 bytes in 0.198 second response time"
"[1680549769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 20 Hours, 57 Minutes"
[1680549863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2295 MiB (1.12% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680549918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.488 second response time
"[1680549961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 11 Hours, 40 Minutes"
[1680550219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.504 second response time
[1680550463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2294 MiB (1.12% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680550519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.706 second response time
"[1680550669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 21 Hours, 12 Minutes"
[1680550764] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2294 MiB (1.12% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.78% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680550818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.684 second response time
"[1680550861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 11 Hours, 55 Minutes"
[1680551063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2294 MiB (1.12% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4583 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680551119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.420 second response time
[1680551363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2262 MiB (1.10% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.77% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680551418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.453 second response time
[1680551468] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680551534] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680551569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 21 Hours, 27 Minutes"
[1680551641] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113670 MB (19.75% inode=100%):
[1680551663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2262 MiB (1.10% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680551663] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114102 MB (19.82% inode=100%):
[1680551683] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs3;WARNING;HARD;5;DISK WARNING - free space: /db2fs3 114102 MB (19.82% inode=100%):
[1680551704] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs4;WARNING;HARD;5;DISK WARNING - free space: /db2fs4 114102 MB (19.82% inode=100%):
[1680551719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.437 second response time
[1680551724] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs5;WARNING;HARD;5;DISK WARNING - free space: /db2fs5 114102 MB (19.82% inode=100%):
"[1680551761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 12 Hours, 10 Minutes"
[1680552018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.480 second response time
[1680552263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2230 MiB (1.08% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680552319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.775 second response time
"[1680552469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 21 Hours, 42 Minutes"
[1680552563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2230 MiB (1.08% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680552618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.400 second response time
"[1680552661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 12 Hours, 25 Minutes"
[1680552919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.638 second response time
[1680553163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2229 MiB (1.08% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680553218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.344 second response time
"[1680553369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 21 Hours, 57 Minutes"
[1680553463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2229 MiB (1.08% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680553519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.486 second response time
"[1680553561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 12 Hours, 40 Minutes"
[1680553689] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680553749] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.058 second response time
[1680553763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2229 MiB (1.08% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680553818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.490 second response time
[1680554063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2229 MiB (1.08% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680554119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.406 second response time
"[1680554269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 22 Hours, 12 Minutes"
[1680554363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2229 MiB (1.08% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4587 MiB (89.76% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680554418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.459 second response time
"[1680554461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 12 Hours, 55 Minutes"
[1680554536] SERVICE ALERT: vswprddb2db11;serGEN vswprd check lock wait;WARNING;SOFT;1;Warning - Warning on lockwait. lockwait is 8
[1680554596] SERVICE ALERT: vswprddb2db11;serGEN vswprd check lock wait;WARNING;HARD;2;Warning - Warning on lockwait. lockwait is 8
[1680554663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2197 MiB (1.07% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6179 MiB (67.77% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680554717] SERVICE ALERT: vswprddb2db11;serGEN vswprd check lock wait;OK;HARD;2;OK - lockwait is 2
[1680554719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.580 second response time
[1680554963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7842 MiB (78.50% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2196 MiB (1.07% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680555018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.428 second response time
"[1680555170] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 22 Hours, 27 Minutes"
[1680555319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.484 second response time
"[1680555361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 13 Hours, 10 Minutes"
[1680555618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.701 second response time
[1680555863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2196 MiB (1.07% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680555919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.719 second response time
"[1680556069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 22 Hours, 42 Minutes"
[1680556163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2196 MiB (1.07% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6197 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680556218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.453 second response time
"[1680556261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 13 Hours, 25 Minutes"
[1680556404] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680556464] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2196 MiB (1.07% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680556474] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1680556519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.431 second response time
"[1680556534] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;HARD;3;HTTP OK: Status line output matched ""200"" - 578260 bytes in 0.097 second response time"
[1680556818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.386 second response time
"[1680556969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 22 Hours, 57 Minutes"
[1680557041] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 110930 MB (19.27% inode=100%):
[1680557063] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 111284 MB (19.33% inode=100%):
[1680557083] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs3;WARNING;HARD;5;DISK WARNING - free space: /db2fs3 111133 MB (19.31% inode=100%):
[1680557104] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs4;WARNING;HARD;5;DISK WARNING - free space: /db2fs4 110897 MB (19.27% inode=100%):
[1680557119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.726 second response time
[1680557124] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs5;WARNING;HARD;5;DISK WARNING - free space: /db2fs5 110776 MB (19.25% inode=100%):
"[1680557161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 13 Hours, 40 Minutes"
[1680557363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2164 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680557419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.520 second response time
[1680557663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2164 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680557719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.240 second response time
"[1680557869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 23 Hours, 12 Minutes"
[1680557941] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 110903 MB (19.27% inode=100%):
[1680557963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2163 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4584 MiB (89.70% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680557964] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 111426 MB (19.36% inode=100%):
[1680557983] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs3;WARNING;HARD;5;DISK WARNING - free space: /db2fs3 111500 MB (19.37% inode=100%):
[1680558004] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs4;WARNING;HARD;5;DISK WARNING - free space: /db2fs4 111671 MB (19.40% inode=100%):
[1680558018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.391 second response time
[1680558024] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs5;WARNING;HARD;5;DISK WARNING - free space: /db2fs5 111709 MB (19.41% inode=100%):
"[1680558061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 13 Hours, 55 Minutes"
[1680558263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2163 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680558319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.600 second response time
[1680558563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2163 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680558618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.625 second response time
"[1680558769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 23 Hours, 27 Minutes"
[1680558841] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 111607 MB (19.39% inode=100%):
[1680558863] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 112041 MB (19.47% inode=100%):
[1680558863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2163 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680558883] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs3;WARNING;HARD;5;DISK WARNING - free space: /db2fs3 112041 MB (19.47% inode=100%):
[1680558904] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs4;WARNING;HARD;5;DISK WARNING - free space: /db2fs4 112042 MB (19.47% inode=100%):
[1680558919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.464 second response time
[1680558924] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs5;WARNING;HARD;5;DISK WARNING - free space: /db2fs5 112040 MB (19.47% inode=100%):
"[1680558961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 14 Hours, 10 Minutes"
[1680559163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2163 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680559218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.567 second response time
[1680559463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2163 MiB (1.05% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680559519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.454 second response time
"[1680559669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 23 Hours, 42 Minutes"
[1680559741] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 111571 MB (19.38% inode=100%):
[1680559763] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 112003 MB (19.46% inode=100%):
[1680559763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2131 MiB (1.04% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680559783] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs3;WARNING;HARD;5;DISK WARNING - free space: /db2fs3 112006 MB (19.46% inode=100%):
[1680559804] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs4;WARNING;HARD;5;DISK WARNING - free space: /db2fs4 114102 MB (19.82% inode=100%):
[1680559819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.548 second response time
[1680559824] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs5;WARNING;HARD;5;DISK WARNING - free space: /db2fs5 114102 MB (19.82% inode=100%):
"[1680559861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 14 Hours, 25 Minutes"
[1680560063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2131 MiB (1.04% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680560118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.492 second response time
[1680560363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2131 MiB (1.04% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680560419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.607 second response time
"[1680560569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 103 Days, 23 Hours, 57 Minutes"
[1680560641] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113670 MB (19.75% inode=100%):
[1680560663] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs2;WARNING;HARD;5;DISK WARNING - free space: /db2fs2 114102 MB (19.82% inode=100%):
[1680560663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2130 MiB (1.04% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680560683] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs3;WARNING;HARD;5;DISK WARNING - free space: /db2fs3 114102 MB (19.82% inode=100%):
[1680560718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.405 second response time
"[1680560762] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 14 Hours, 40 Minutes"
[1680561019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.476 second response time
[1680561263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2130 MiB (1.04% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680561318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.543 second response time
"[1680561469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 12 Hours, 12 Minutes"
[1680561563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2130 MiB (1.04% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4580 MiB (89.63% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680561619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.796 second response time
"[1680561661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 14 Hours, 55 Minutes"
[1680561864] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2098 MiB (1.02% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680561918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.411 second response time
[1680562163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2098 MiB (1.02% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680562219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.588 second response time
"[1680562369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 27 Hours, 27 Minutes"
[1680562463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2098 MiB (1.02% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4576 MiB (89.55% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680562518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.413 second response time
"[1680562561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 15 Hours, 10 Minutes"
[1680562763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2098 MiB (1.02% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680562819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.601 second response time
[1680563063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2098 MiB (1.02% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680563118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.527 second response time
"[1680563269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 42 Hours, 42 Minutes"
[1680563363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2097 MiB (1.02% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680563419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.554 second response time
"[1680563461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 15 Hours, 25 Minutes"
[1680563718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.392 second response time
[1680563878] SERVICE ALERT: svpdrpau33;serGEN vswprd Drupal HTTP Connect;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 200 OK - 577605 bytes in 4.549 second response time
[1680563878] SERVICE ALERT: svpdrpau32;serGEN vswprd Drupal HTTP Connect 443;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 200 OK - 577616 bytes in 5.137 second response time
[1680563938] SERVICE ALERT: svpdrpau33;serGEN vswprd Drupal HTTP Connect;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 577604 bytes in 0.057 second response time
[1680563939] SERVICE ALERT: svpdrpau32;serGEN vswprd Drupal HTTP Connect 443;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 577615 bytes in 0.061 second response time
[1680563963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2097 MiB (1.02% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680564019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.443 second response time
"[1680564169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 57 Hours, 57 Minutes"
[1680564263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2065 MiB (1.00% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.96% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680564318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.479 second response time
"[1680564361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 15 Hours, 40 Minutes"
[1680564619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.429 second response time
[1680564628] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680564688] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.069 second response time
[1680564863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2033 MiB (0.99% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4579 MiB (89.62% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680564918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.473 second response time
"[1680565069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 1 Hours, 12 Minutes"
"[1680565143] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;WARNING;SOFT;1;HTTP WARNING: Status line output matched ""200"" - 578260 bytes in 7.368 second response time"
"[1680565144] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;WARNING;SOFT;1;HTTP WARNING: Status line output matched ""200"" - 578260 bytes in 7.366 second response time"
[1680565163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2033 MiB (0.99% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
"[1680565203] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578260 bytes in 0.317 second response time"
"[1680565205] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578260 bytes in 0.316 second response time"
[1680565219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.511 second response time
"[1680565261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 15 Hours, 55 Minutes"
[1680565518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.497 second response time
[1680565544] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680565546] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680565604] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578259 bytes in 0.100 second response time"
"[1680565607] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578223 bytes in 0.773 second response time"
[1680565763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2033 MiB (0.99% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680565819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.477 second response time
"[1680565969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 1 Hours, 27 Minutes"
[1680566063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2033 MiB (0.99% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680566118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.502 second response time
"[1680566161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 16 Hours, 10 Minutes"
[1680566363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2032 MiB (0.99% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680566419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.748 second response time
[1680566663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2000 MiB (0.97% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680566718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.483 second response time
"[1680566869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 1 Hours, 42 Minutes"
[1680566963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2000 MiB (0.97% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680567019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.517 second response time
"[1680567061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 16 Hours, 25 Minutes"
[1680567263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2000 MiB (0.97% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515890 MiB (53.31% inode=100%):
[1680567319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.483 second response time
[1680567474] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680567534] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578260 bytes in 0.095 second response time"
[1680567618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.636 second response time
"[1680567747] SERVICE ALERT: vswprdjboss94;serGEN vswprd Stratus Authorization status;CRITICAL;SOFT;1;Critical - stratus_authorization returned False, Last authorization sent @ 2023-04-04T00:16:49.371Z: last heartbeat sent @ 2023-04-04T00:20:00.090Z"
"[1680567765] SERVICE ALERT: vswprdjboss93;serGEN vswprd Stratus OnlineSocket Status;WARNING;SOFT;1;Warning - stratus_online_socket returned False, Last heartbeat response was 166 seconds ago (request sequence 3508)"
"[1680567769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 1 Hours, 57 Minutes"
"[1680567783] SERVICE ALERT: vswprdjboss94;serGEN vswprd Stratus OnlineSocket Status;WARNING;SOFT;1;Warning - stratus_online_socket returned False, Last heartbeat response was 184 seconds ago (request sequence 2025)"
[1680567807] SERVICE ALERT: vswprdjboss94;serGEN vswprd Stratus Authorization status;OK;SOFT;2;OK - Last authorization successfully sent @ 2023-04-04T00:22:56.357Z
"[1680567825] SERVICE ALERT: vswprdjboss93;serGEN vswprd Stratus OnlineSocket Status;WARNING;SOFT;2;Warning - stratus_online_socket returned False, Last heartbeat response was 226 seconds ago (request sequence 3508)"
[1680567843] SERVICE ALERT: vswprdjboss94;serGEN vswprd Stratus OnlineSocket Status;OK;SOFT;2;OK - Last heartbeat response was 4 seconds ago (sequence=2027: host=206.253.180.115: port=4536)
[1680567863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2000 MiB (0.97% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680567885] SERVICE ALERT: vswprdjboss93;serGEN vswprd Stratus OnlineSocket Status;OK;HARD;3;OK - Last heartbeat response was 46 seconds ago (sequence=3510: host=206.253.180.115: port=4536)
[1680567919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.539 second response time
"[1680567961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 16 Hours, 40 Minutes"
[1680568093] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680568154] SERVICE ALERT: vswprdwpxwa01;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578283 bytes in 1.090 second response time"
[1680568163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 2000 MiB (0.97% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680568218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.769 second response time
[1680568519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.717 second response time
"[1680568670] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 2 Hours, 12 Minutes"
[1680568763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1968 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6178 MiB (67.76% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680568819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.619 second response time
"[1680568861] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 16 Hours, 55 Minutes"
[1680568977] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680569037] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578258 bytes in 0.368 second response time"
[1680569063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1967 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680569118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.575 second response time
"[1680569404] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;WARNING;SOFT;1;HTTP WARNING: Status line output matched ""200"" - 578236 bytes in 7.383 second response time"
[1680569404] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680569419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 1.003 second response time
"[1680569464] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578259 bytes in 0.120 second response time"
"[1680569464] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578319 bytes in 0.075 second response time"
"[1680569569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 2 Hours, 27 Minutes"
[1680569663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1967 MiB (0.96% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680569719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.367 second response time
"[1680569761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 17 Hours, 10 Minutes"
[1680570018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.596 second response time
[1680570319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.309 second response time
"[1680570469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 2 Hours, 42 Minutes"
[1680570618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.573 second response time
"[1680570661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 17 Hours, 25 Minutes"
[1680570863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1935 MiB (0.94% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680570919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.962 second response time
[1680571163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1935 MiB (0.94% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680571219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.754 second response time
[1680571271] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave11;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
"[1680571369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 2 Hours, 57 Minutes"
[1680571389] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_BulkImports;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571451] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave11;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.593 second response time
"[1680571562] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 17 Hours, 40 Minutes"
[1680571563] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_BuildBidPacket;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571569] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_BulkImports;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571582] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ExcelExport;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571604] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromMerlot;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571623] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromSolver;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571632] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave11;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680571635] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_StartDailyOtaAccountingImports;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571644] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ProcessingStatus;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571745] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_BuildBidPacket;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571750] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_BulkImports;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680571763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1934 MiB (0.94% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680571764] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ExcelExport;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571787] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromMerlot;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571806] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromSolver;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571812] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave11;UNKNOWN;SOFT;4;Unknown - Query failed with status code: 401
[1680571816] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_StartDailyOtaAccountingImports;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.386 second response time
[1680571826] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ProcessingStatus;UNKNOWN;SOFT;2;Unknown - Query failed with status code: 401
[1680571930] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_BuildBidPacket;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680571931] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_BulkImports;UNKNOWN;SOFT;4;Unknown - Query failed with status code: 401
[1680571934] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680571934] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680571946] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ExcelExport;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680571969] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromMerlot;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680571978] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave12;UNKNOWN;SOFT;1;Unknown - Query failed with status code: 401
[1680571988] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromSolver;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680571993] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave11;UNKNOWN;HARD;5;Unknown - Query failed with status code: 401
"[1680571994] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578235 bytes in 0.344 second response time"
"[1680571995] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578319 bytes in 0.355 second response time"
[1680571996] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_StartDailyOtaAccountingImports;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680572009] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ProcessingStatus;UNKNOWN;SOFT;3;Unknown - Query failed with status code: 401
[1680572063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1934 MiB (0.94% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.28% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680572111] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_BuildBidPacket;OK;SOFT;4;OK - 2023-04-03_09-17-03
[1680572115] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_BulkImports;OK;HARD;5;OK - BulkImports OK
[1680572118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.304 second response time
[1680572127] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ExcelExport;OK;SOFT;4;OK - 2023-04-03_11-44-57
[1680572149] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromMerlot;OK;SOFT;4;OK - 2023-04-03_10-15-57
[1680572163] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave12;OK;SOFT;2;OK - vswprdjnkap11_slave12 OK
[1680572169] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ImportFromSolver;OK;SOFT;4;OK - 2023-03-21_13-49-17
[1680572181] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_StartDailyOtaAccountingImports;OK;SOFT;4;OK - StartDailyOtaAccountingImports OK
[1680572189] SERVICE ALERT: cbi-admin;serGEN vswprd jenkins_CBI_ProcessingStatus;OK;SOFT;4;OK - 2023-04-03_08-52-31
"[1680572269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 3 Hours, 12 Minutes"
[1680572419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.521 second response time
"[1680572461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 17 Hours, 55 Minutes"
[1680572664] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1934 MiB (0.94% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4569 MiB (89.41% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680572718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.372 second response time
[1680572898] SERVICE ALERT: vswprdjnkap11;serGEN vswprd jenkins_vswprdjnkap11_slave11;OK;HARD;5;OK - vswprdjnkap11_slave11 OK
[1680572963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1934 MiB (0.94% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4556 MiB (89.16% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680573019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.479 second response time
"[1680573169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 3 Hours, 27 Minutes"
[1680573263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1902 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6196 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4556 MiB (89.16% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680573318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.425 second response time
"[1680573361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 18 Hours, 10 Minutes"
[1680573563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1902 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4556 MiB (89.16% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680573619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.402 second response time
[1680573918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.479 second response time
"[1680574069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 3 Hours, 42 Minutes"
[1680574219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.481 second response time
"[1680574262] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 18 Hours, 25 Minutes"
[1680574463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1902 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4556 MiB (89.16% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680574518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.503 second response time
[1680574763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1901 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4556 MiB (89.16% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680574819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.523 second response time
"[1680574969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 3 Hours, 57 Minutes"
[1680575063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1901 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4556 MiB (89.16% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680575118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.564 second response time
"[1680575161] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 18 Hours, 40 Minutes"
[1680575363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1901 MiB (0.92% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680575418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.291 second response time
[1680575567] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680575627] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.067 second response time
[1680575663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1869 MiB (0.91% inode=100%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.94% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680575719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.488 second response time
"[1680575869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 4 Hours, 12 Minutes"
[1680575963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 305 MiB (0.14% inode=97%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6195 MiB (67.95% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4571 MiB (89.46% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680576018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.587 second response time
"[1680576061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 18 Hours, 55 Minutes"
[1680576263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1517 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6177 MiB (67.75% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680576318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.242 second response time
[1680576563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1516 MiB (0.74% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680576618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.631 second response time
"[1680576769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 4 Hours, 27 Minutes"
[1680576863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1484 MiB (0.72% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4568 MiB (89.40% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680576919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.465 second response time
"[1680576961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 19 Hours, 10 Minutes"
[1680577219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.655 second response time
[1680577273] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680577275] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680577333] SERVICE ALERT: vswprdwpxwa04;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578319 bytes in 0.242 second response time"
"[1680577335] SERVICE ALERT: vswprdwpxwa06;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578295 bytes in 0.235 second response time"
[1680577464] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680577465] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680577518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.385 second response time
"[1680577524] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578222 bytes in 0.267 second response time"
"[1680577525] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578297 bytes in 0.250 second response time"
[1680577594] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680577595] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;CRITICAL;SOFT;1;CRITICAL - Socket timeout
"[1680577654] SERVICE ALERT: vswprdwpxwa03;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578319 bytes in 0.110 second response time"
"[1680577656] SERVICE ALERT: vswprdwpxwa05;serGEN vswprd www rebrand proxy check;OK;SOFT;2;HTTP OK: Status line output matched ""200"" - 578223 bytes in 0.688 second response time"
"[1680577669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 4 Hours, 42 Minutes"
[1680577763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1452 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680577818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.512 second response time
"[1680577862] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 19 Hours, 25 Minutes"
[1680578119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.577 second response time
[1680578419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.569 second response time
"[1680578569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 4 Hours, 57 Minutes"
[1680578663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1452 MiB (0.70% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680578718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.473 second response time
"[1680578761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 19 Hours, 40 Minutes"
[1680578791] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;WARNING;SOFT;1;WARNING - 361 established connections to port 50001
[1680578851] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;OK;SOFT;2;OK - 358 established connections to port 50001
[1680578963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1419 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680579019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.525 second response time
[1680579096] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680579162] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680579316] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680579318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.474 second response time
[1680579382] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680579426] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680579469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 5 Hours, 12 Minutes"
[1680579491] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680579619] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.462 second response time
"[1680579661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 19 Hours, 55 Minutes"
[1680579863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1419 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680579918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.589 second response time
[1680580163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1419 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680580219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.430 second response time
"[1680580369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 5 Hours, 27 Minutes"
[1680580463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1419 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
"[1680580562] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 20 Hours, 10 Minutes"
[1680580763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1419 MiB (0.69% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680580819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.647 second response time
[1680581063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1387 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.88% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680581118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.185 second response time
"[1680581270] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 5 Hours, 42 Minutes"
[1680581363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1387 MiB (0.67% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680581394] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680581418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.477 second response time
[1680581459] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
"[1680581461] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 20 Hours, 25 Minutes"
[1680581524] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680581719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.497 second response time
"[1680581940] SERVICE ALERT: vswprdaisap02;serGEN vswprd AIS CMSARINC Server;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with args 'com.cms_solutions.cmsarincserver.CMSARINCServer', UID = 0 (root)"
[1680581963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1396 MiB (0.68% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
"[1680582000] SERVICE ALERT: vswprdaisap02;serGEN vswprd AIS CMSARINC Server;OK;SOFT;2;PROCS OK: 1 process with args 'com.cms_solutions.cmsarincserver.CMSARINCServer', UID = 0 (root)"
[1680582018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.385 second response time
"[1680582169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 5 Hours, 57 Minutes"
[1680582263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1396 MiB (0.68% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680582319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.483 second response time
"[1680582361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 20 Hours, 40 Minutes"
[1680582563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1396 MiB (0.68% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680582618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.414 second response time
[1680582667] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680582733] SERVICE ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680582811] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;WARNING;SOFT;1;WARNING - 361 established connections to port 50001
[1680582828] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680582863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1396 MiB (0.68% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680582871] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;WARNING;SOFT;2;WARNING - 361 established connections to port 50001
[1680582893] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680582919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.513 second response time
[1680582931] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;WARNING;HARD;3;WARNING - 361 established connections to port 50001
[1680582958] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680583023] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680583069] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 6 Hours, 12 Minutes"
[1680583163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1364 MiB (0.66% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6177 MiB (67.74% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4565 MiB (89.34% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680583218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.493 second response time
[1680583231] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;OK;HARD;3;OK - 358 established connections to port 50001
"[1680583261] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 20 Hours, 55 Minutes"
[1680583463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1364 MiB (0.66% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680583519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.474 second response time
[1680583641] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680583707] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680583730] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680583764] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1364 MiB (0.66% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680583795] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680583818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.490 second response time
[1680583848] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680583913] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680583969] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 6 Hours, 27 Minutes"
[1680584063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1364 MiB (0.66% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680584073] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680584119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.458 second response time
[1680584138] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
"[1680584162] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 21 Hours, 10 Minutes"
[1680584203] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680584299] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680584363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1364 MiB (0.66% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4595 MiB (89.93% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680584363] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680584364] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680584418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.700 second response time
[1680584419] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680584429] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680584430] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680584485] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680584663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1363 MiB (0.66% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4595 MiB (89.93% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680584719] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.443 second response time
[1680584735] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680584869] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 6 Hours, 42 Minutes"
[1680584948] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680585013] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680585018] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.273 second response time
"[1680585061] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 21 Hours, 25 Minutes"
[1680585109] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680585174] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680585215] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680585257] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680585263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1331 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6189 MiB (67.87% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.95% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680585280] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680585319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.515 second response time
[1680585322] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680585340] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680585405] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680585563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1331 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.94% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680585618] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.499 second response time
[1680585645] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680585661] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680585710] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680585726] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
"[1680585769] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 6 Hours, 57 Minutes"
[1680585792] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680585863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1331 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6194 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.94% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680585918] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.406 second response time
"[1680585961] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 21 Hours, 40 Minutes"
[1680586029] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680586075] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680586079] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680586094] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680586139] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680586140] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680586144] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680586159] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680586163] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1331 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.94% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680586204] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680586206] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680586219] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.440 second response time
[1680586270] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680586308] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680586350] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680586373] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680586415] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680586439] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680586463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1331 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4596 MiB (89.94% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680586506] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1680586511] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680586518] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.469 second response time
[1680586566] SERVICE ALERT: vswprdloyad11;serGEN vswprd Loyalty administration site;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 8259 bytes in 0.072 second response time
"[1680586669] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 7 Hours, 12 Minutes"
[1680586763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1331 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4596 MiB (89.94% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680586819] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.546 second response time
[1680586819] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680586831] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;WARNING;SOFT;1;WARNING - 361 established connections to port 50001
"[1680586862] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 21 Hours, 55 Minutes"
[1680586884] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680586891] SERVICE ALERT: vswprdawbws12;serGEN vswprd Airweb - db2 established connections;OK;SOFT;2;OK - 358 established connections to port 50001
[1680587063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1331 MiB (0.65% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6192 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680587119] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.587 second response time
[1680587363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1330 MiB (0.64% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6192 MiB (67.91% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680587381] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680587409] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680587416] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680587418] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.238 second response time
[1680587446] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680587475] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680587481] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680587511] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680587543] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
"[1680587569] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 7 Hours, 27 Minutes"
[1680587577] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680587609] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680587641] SERVICE ALERT: vswprddb2db11;serGEN vswprd Volume - db2fs1;WARNING;HARD;5;DISK WARNING - free space: /db2fs1 113669 MB (19.75% inode=100%):
[1680587642] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680587663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1298 MiB (0.63% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2515889 MiB (53.31% inode=100%):
[1680587718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.544 second response time
"[1680587761] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 22 Hours, 10 Minutes"
[1680587850] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680587915] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680587964] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680588019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.839 second response time
[1680588030] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680588096] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680588171] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680588236] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680588263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1298 MiB (0.63% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2513646 MiB (53.27% inode=100%):
[1680588274] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680588319] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.456 second response time
[1680588339] SERVICE ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680588469] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 7 Hours, 42 Minutes"
[1680588547] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680588563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1298 MiB (0.63% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.93% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2507874 MiB (53.14% inode=100%):
[1680588612] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680588661] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 22 Hours, 25 Minutes"
[1680588863] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1298 MiB (0.63% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2502145 MiB (53.02% inode=100%):
[1680588919] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.433 second response time
[1680588976] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680588979] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589041] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680589044] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680589109] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589164] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1298 MiB (0.63% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2496508 MiB (52.90% inode=100%):
[1680589175] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680589218] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.415 second response time
[1680589240] SERVICE ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680589369] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 7 Hours, 57 Minutes"
[1680589398] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589463] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1298 MiB (0.63% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2490775 MiB (52.78% inode=100%):
[1680589463] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680589519] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.515 second response time
"[1680589561] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 22 Hours, 40 Minutes"
[1680589641] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589647] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589688] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589706] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680589713] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680589754] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680589763] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1234 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2485175 MiB (52.66% inode=100%):
[1680589771] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680589778] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680589781] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589802] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589802] SERVICE DOWNTIME ALERT: www-status;serGEN vswprd F5 Sessions;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap13;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap16;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap27;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589802] SERVICE DOWNTIME ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;STARTED; Service has entered a period of scheduled downtime
[1680589818] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.224 second response time
[1680589846] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680589867] SERVICE ALERT: vswprdgqlap21;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680589895] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589911] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680589946] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680589960] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590011] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590063] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1234 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2479496 MiB (52.54% inode=100%):
[1680590076] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680590118] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.466 second response time
[1680590177] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590187] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590201] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590217] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590242] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590253] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680590267] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
"[1680590269] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 8 Hours, 12 Minutes"
[1680590282] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680590307] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590318] SERVICE ALERT: vswprdgqlap12;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680590325] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590332] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;WARNING;HARD;3;Warning - Warning on count. snmp query result
[1680590345] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590348] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680590363] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1233 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4592 MiB (89.88% inode=100%): /nvol_prd_db2_bku_01 2473692 MiB (52.42% inode=100%):
[1680590373] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590382] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590383] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590390] SERVICE ALERT: vswprdgqlap14;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590410] SERVICE ALERT: vswprdgqlap28;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590419] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.476 second response time
[1680590447] SERVICE ALERT: vswprdgqlap15;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590449] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
"[1680590462] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 22 Hours, 55 Minutes"
[1680590638] SERVICE ALERT: vswprdgqlap22;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680590663] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1233 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10156 MiB (99.28% inode=100%): /var/log/audit 4588 MiB (89.80% inode=100%): /nvol_prd_db2_bku_01 2468140 MiB (52.30% inode=100%):
[1680590718] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.497 second response time
[1680590729] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590794] SERVICE ALERT: vswprdgqlap25;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590814] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590880] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680590945] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590953] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680590963] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1233 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6191 MiB (67.90% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10157 MiB (99.29% inode=100%): /var/log/audit 4588 MiB (89.80% inode=100%): /nvol_prd_db2_bku_01 2424412 MiB (51.38% inode=100%):
[1680590976] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680591010] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680591019] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.498 second response time
[1680591019] SERVICE ALERT: vswprdgqlap23;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680591039] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680591042] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680591075] SERVICE ALERT: vswprdgqlap17;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680591104] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680591107] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680591169] SERVICE ALERT: vswprdrhcap01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 104 Days, 8 Hours, 27 Minutes"
[1680591169] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680591208] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680591235] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680591263] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1233 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4588 MiB (89.80% inode=100%): /nvol_prd_db2_bku_01 2369377 MiB (50.21% inode=100%):
[1680591273] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;WARNING;SOFT;2;Warning - Warning on count. snmp query result
[1680591300] SERVICE ALERT: vswprdgqlap11;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
[1680591318] SERVICE ALERT: flight;serGEN vswprd Allegiant flight deals flight.deals.allegiant.com;CRITICAL;HARD;3;HTTP CRITICAL: HTTP/1.1 302 Moved Temporarily - string 'Allegiant' not found on 'https://flight.deals.allegiant.com:3128https://flight.deals.allegiant.com/ats/url.aspx?cr=986&wu=95' - 793 bytes in 0.448 second response time
[1680591319] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680591338] SERVICE ALERT: vswprdgqlap18;serGEN vswprd f5 connection counts;OK;HARD;3;OK - snmp query result
"[1680591361] SERVICE ALERT: svpjmplx01;serGEN vswprd RHEL Uptime;WARNING;HARD;5;WARNING - system has not rebooted for 369 Days, 23 Hours, 10 Minutes"
[1680591385] SERVICE ALERT: vswprdgqlap24;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680591412] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680591478] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;OK;SOFT;2;OK - snmp query result
[1680591543] SERVICE ALERT: vswprdgqlap26;serGEN vswprd f5 connection counts;WARNING;SOFT;1;Warning - Warning on count. snmp query result
[1680591563] SERVICE ALERT: vswprddsmap01;serGEN vswprd All Volumes;WARNING;HARD;3;DISK WARNING - free space: / 7971 MiB (79.79% inode=98%): /boot 268 MiB (53.99% inode=100%): /data1 1233 MiB (0.60% inode=99%): /tmp 5077 MiB (99.36% inode=100%): /var 13624 MiB (88.76% inode=100%): /opt 6193 MiB (67.92% inode=100%): /var/tmp 5077 MiB (99.36% inode=100%): /var/log 10158 MiB (99.29% inode=100%): /var/log/audit 4588 MiB (89.80% inode=100%): /nvol_prd_db2_bku_01 2314966 MiB (49.06% inode=100%):